Syngap1 regulates the synaptic drive and membrane excitability of Parvalbumin-positive interneurons in mouse auditory cortex

Author Response:

The following is the authors’ response to the previous reviews

Reviewer #2 (Public review):

Summary:

In this manuscript, the authors investigated how partial loss of SynGap1 affects inhibitory neurons derived from the MGE in the auditory cortex, focusing on their synaptic inputs and excitability. While haplo-insufficiently of SynGap1 is known to lead to intellectual disabilities, the underlying mechanisms remain unclear.

Strengths:

The questions are novel

Weaknesses:

Despite the interesting and novel questions, there are significant issues regarding the experimental design and potential misinterpretations of key findings. Consequently, the manuscript contributes little to our understanding of SynGap1 loss mechanisms.

Major issues in the second version of the manuscript:

In the review of the first version there were major issues and contradictions with the sEPSC and mEPSC data, and were not resolved after the revision, and the new control experiments rather confirmed the contradiction.

In the original review I stated: "One major concern is the inconsistency and confusion in the intermediate conclusions drawn from the results. For instance, while the sEPSC data indicates decreased amplitude in PV+ and SOM+ cells in cHet animals, the frequency of events remains unchanged. In contrast, the mEPSC data shows no change in amplitudes in PV+ cells, but a significant decrease in event frequency. The authors conclude that the former observation implies decreased excitability. However, traditionally, such observations on mEPSC parameters are considered indicative of presynaptic mechanisms rather than changes of network activity. The subsequent synapse counting experiments align more closely with the traditional conclusions. This issue can be resolved by rephrasing the text. However, it would remain unexplained why the sEPSC frequency shows no significant difference. If the majority of sEPSC events were indeed mediated by spiking (which is blocked by TTX), the average amplitudes and frequency of mEPSCs should be substantially lower than those of sEPSCs. Yet, they fall within a very similar range, suggesting that most sEPSCs may actually be independent of action potentials. But if that was indeed the case, the changes of purported sEPSC and mEPSC results should have been similar." Contradictions remained after the revision of the manuscript. On one hand, the authors claimed in the revised version that "We found no difference in mEPSC amplitude between the two genotypes (Fig. 1g), indicating that the observed difference in sEPSC amplitude (Figure 1b) could arise from decreased network excitability". On the other hand, later they show "no significative difference in either amplitude or inter-event intervals between sEPSC and mEPSC, suggesting that in acute slices from adult A1, most sEPSCs may actually be AP independent." The latter means that sEPSCs and mEPSCs are the same type of events, which should have the same sensitivity to manipulations.

We thank the reviewer for the detailed comments. Our results suggest a diverse population of PV+ cells, with varying reliance on action potential-dependent and -independent release. Several PV+ cells indeed show TTX sensitivity (reduced EPSC event amplitudes following TTX application: See new Supplementary Figure 2b-e), but their individual responses are diluted when all cells are pooled together. To account for this variability, we recorded sEPSC followed by mEPSC from more mice of both genotypes (new Figure 1f-j). Further, following the editors and reviewers’ suggestions, we removed speculations about the role of network activity changes.

In summary, our data confirmed that TTX blocked APs in PV+ cells and that recordings were stable as indicated by lack of changes in series resistance during the recording period in our experimental setup (new Suppl. Figure 2f-i). We found no difference in mEPSC amplitude between the two genotypes (Fig. 1g, right), indicating that the observed difference in sEPSC amplitude (Figure 1c, right) could be due to impaired AP-dependent release in cHet mice and the presence of large-amplitude sEPSCs that are preferentially affected by TTX in control mice (new Suppl. Figure 2b-e). Conversely, cHet mice showed longer inter-mEPSC time interval (cumulative distribution in Figure 1g, left), and significantly lower charge transfer and DQ*f (Figure 1j) compared to controls littermates, suggesting a decrease of glutamatergic presynaptic release sites onto PV+ cells.

Concerns about the quality of the synapse counting experiments were addressed by showing additional images in a different and explaining quantification. However, the admitted restriction of the analysis of excitatory synapses to the somatic region represent a limitation, as they include only a small fraction of the total excitation - even if, the slightly larger amplitudes of their EPSPs are considered.

We agree with the reviewer that restricting the anatomical analysis of excitatory synapses to PV cell somatic region is a limitation, as highlighted it in the discussion of the revised manuscript. Recent studies, based on serial block-face scanning electron microscopy, suggest that cortical PV+ interneurons receive more robust excitatory inputs to their perisomatic region as compared to pyramidal neurons (see for example, Hwang et al. 2021, Cerebral Cortex, http://doi.org/10.1093/cercor/bhaa378). It is thus possible that putative glutamatergic synapses, analysed by vGlut1/PSD95 colocalisation around PV+ cell somata, may be representative of a substantially major excitatory input population. Since analysing putative excitatory synapses onto PV+ dendrites would be difficult and require a much longer time, we re-phrased the text to more clearly highlight the rationale and limitation of this approach.

New experiments using paired-pulse stimulation provided an answer to issues 3 and 4. Note that the numbering of the Figures in the responses and manuscript are not consistent.

We are glad that the reviewer found that the new paired-pulse experiments answered previously raised concerns. We corrected the discrepancy in figure numbers in the manuscript. Thank you for noticing.

I agree that low sampling rate of the APs does not change the observed large differences in AP threshold, however, the phase plots are still inconsistent in a sense that there appears to be an offset, as all values are shifted to more depolarized membrane potentials, including threshold, AP peak, AHP peak. This consistent shift may be due to a non-biological differences in the two sets of recordings, and, importantly, it may negate the interpretation of the I/f curves results (Fig. 5e).

We agree with the reviewers that higher sampling rate would allow to more accurately assess different parameters, such as AP peak, half-width, rise time, etc., while it would not affect the large differences in AP threshold we observed between control and mutant mice. Since the phase plots to not add to our result analysis, we removed them from the revised manuscript.

Additional issues:

The first paragraph of the Results mentioned that the recorded cells were identified by immunolabelling and axonal localization. However, neither the Results nor the Methods mention the criteria and levels of measurements of axonal arborization.

Recorded MGE-derived interneurons were filled with biocytin, and their identity was confirmed by immunolabeling for neurochemical markers (PV or SST) and analysis of anatomical properties. In particular, whole biocytin-positive immunolabelled neurons were acquired using a Leica SP8-DLS confocal microscope (20x objective, NA 0.75; Z-step 1 1μm). For each imaged neuron, which was the result of multiple merged confocal stacks, we visually determined the spatial distribution across cortical layers of the axonal arbor and whether its dendrites carried spines. We added this information in the method section. Furthermore, to better represent our methodological approach, we added a new figure (Supplemental Figure 1) including 1) two examples of PV+ interneurons, showing dendrites devoid of spines and axons spreading from Layer II to Layer V (new Suppl. Figure 1a); and 2) two examples of SST+ interneurons showing dendritic with spines and axons projecting from Layer IV to Layer I where they gave rise to multiple collaterals (new Suppl. Figure 1b).

The other issues of the first review were adequately addressed by the Authors and the manuscript improved by these changes.

We are happy the reviewer found that the other issues were well addressed.

Reviewer #3 (Public review):

This paper compares the synaptic and membrane properties of two main subtypes of interneurons (PV+, SST+) in the auditory cortex of control mice vs mutants with Syngap1 haploinsufficiency. The authors find differences between control and mutants in both interneuron populations, although they claim a predominance in PV+ cells. These results suggest that altered PVinterneuron functions in the auditory cortex may contribute to the network dysfunctions observed in Syngap1 haploinsufficiency-related intellectual disability.

The subject of the work is interesting, and most of the approach is rather direct and straightforward, which are strengths. There are also some methodological weaknesses and interpretative issues that reduce the impact of the paper.

(1) Supplementary Figure 3: recording and data analysis. The data of Supplementary Figure 3 show no differences either in the frequency or amplitude of synaptic events recorded from the same cell in control (sEPSCs) vs TTX (mEPSCs). This suggests that, under the experimental conditions of the paper, sEPSCs are AP-independent quantal events. However, I am concerned by the high variability of the individual results included in the Figure. Indeed, several datapoints show dramatically different frequencies in control vs TTX, which may be explained by unstable recording conditions. It would be important to present these data as time course plots, so that stability can be evaluated. Also, the claim of lack of effect of TTX should be corroborated by positive control experiments verifying that TTX is working (block of action potentials, for example). Lastly, it is not clear whether the application of TTX was consistent in time and duration in all the experiments and the paper does not clarify what time window was used for quantification.

We understand the reviewer’s concern about high variability. To account for this variability, we recorded sEPSC followed by mEPSC from more mice of both genotypes (see new Figure 1f-j). We confirmed that TTX worked as expected several times through the time course of this study, in different aliquots prepared from the same TTX vial that was used for all experiments. The results of the last test we performed, showing that TTX application blocks action potentials in a PV+ cell, are depicted in new Suppl. Figure 2a. Furthermore, new Suppl. Figure 2f-i shows series resistance (Rs) over time for 4 different PV+ interneurons, indicating recording stability. These results are representative of the entire population of recorded neurons, which we have meticulously analysed one by one. TTX was applied using the same protocol for all recorded neurons. In particular, sEPSCs were first sampled over a 2 min period. A TTX (1μM; Alomone Labs)-containing solution was then perfused into the recording chamber at a flow rate of 2 mL/min. We then waited for 5 min before sampling mEPSCs over a 2 min period. We added this information in the revised manuscript methods.

(2) Figure 1 and Supplementary Figure 3: apparent inconsistency. If, as the authors claim, TTX does not affect sEPSCs (either in the control or mutant genotype, Supplementary Figure 3 and point 1 above), then comparing sEPSC and mEPSC in control vs mutants should yield identical results. In contrast, Figure 1 reports a _selective_ reduction of sEPSCs amplitude (not in mEPSCs) in mutants, which is difficult to understand. The proposed explanation relying on different pools of synaptic vesicles mediating sEPSCs and mEPSCs does not clarify things. If this was the case, wouldn't it also imply a decrease of event frequency following TTX addition? However, this is not observed in Supplementary Figure 3. My understanding is that, according to this explanation, recordings in control solution would reflect the impact of two separate pools of vesicles, whereas, in the presence of TTX, only one pool would be available for release. Therefore, TTX should cause a decrease in the frequency of the recorded events, which is not what is observed in Supplementary Figure 3.

To account for the large variability and clarify these results, we recorded sEPSCs followed by mEPSCs from more mice of both genotypes (new Figure 1f-j). We found no difference in mEPSC amplitude between the two genotypes (Fig. 1g, right), indicating that the observed difference in sEPSC amplitude (Figure 1c, right) could be due to impaired AP-dependent release in cHet mice and the presence of large-amplitude sEPSCs that are preferentially affected by TTX in control mice (new Suppl. Figure 2b-e). Conversely, cHet mice showed longer inter-mEPSC time interval (cumulative distribution in Figure 1g, left), and significantly lower charge transfer and DQ*f (Figure 1j) compared to controls littermates, suggesting a decrease of glutamatergic presynaptic release sites. We rephrased the text in the revised manuscript according to the updated data and, following the reviewer’s suggestions, we removed speculations relying on different pools of synaptic vesicles.

(3) Figure 1: statistical analysis. Although I do appreciate the efforts of the authors to illustrate both cumulative distributions and plunger plots with individual data, I am confused by how the cumulative distributions of Figure 1b (sEPSC amplitude) may support statistically significant differences between genotypes, but this is not the case for the cumulative distributions of Figure 1g (inter mEPSC interval), where the curves appear even more separated. A difference in mEPSC frequency would also be consistent with the data of Supplementary Fig 2b, which otherwise are difficult to reconciliate. I would encourage the authors to use the Kolmogorov-Smirnov rather than a t-test for the comparison of cumulative distributions.

We thank the reviewer for this thoughtful suggestion. We recorded more mice of both genotypes and the updated data now show a significant difference between the cumulative distributions of the inter mEPSC intervals recorded from the two genotypes (new Figure 1g). For statistical analysis, we based our conclusion on the statistical results generated by LMM, modelling animal as a random effect and genotype as fixed effect. We used this statistical analysis because we considered the number of mice as independent replicates and the number of cells in each mouse as repeated measures (Berryer et al. 2016; Heggland et al., 2019; Yu et al., 2022). For cumulative distributions, the same number of events was chosen randomly from each cell and analysed by LMM, modelling animal as a random effect and genotype as fixed effect. The reason we decided to use LMM for our statistical analyses is based on the growing concern over reproducibility in biomedical research and the ongoing discussion on how data are analysed (see for example, Yu et al (2022), Neuron 110:21-35 https://doi.org/10.1016/j.neuron.2021.10.030; Aarts et al. (2014). Nat Neurosci 17, 491–496. https://doi.org/10.1038/nn.3648). We acknowledge that patch-clamp data has been historically analysed using t-test and analysis of variance (ANOVA), or equivalent nonparametric tests. However, these tests assume that individual observations (recorded neurons in this case) are independent of each other. Whether neurons from the same mouse are independent or correlated variables is an unresolved question, but does not appear to be likely from a biological point of view. Statisticians have developed effective methods to analyze correlated data, including LMM.

(4) Methods. I still maintain that a threshold at around -20/-15 mV for the first action potential of a train seems too depolarized (see some datapoints of Fig 5c and Fig7c) for a healthy spike. This suggest that some cells were either in precarious conditions or that the capacitance of the electrode was not compensated properly.

As suggested by the reviewer, in the revised figures we excluded the neurons with threshold at -20/-15 mV. In addition, we performed statistical analysis with and without these cells (data reported below) and found that whether these cells are included or excluded, the statistical significance of the results does not change.

Fig.5c: including the 2 outliers from cHet group with values of -16.5 and 20.6 mV: 42.6±1.01 mV in control, n=33 cells from 15 mice vs -35.3±1.2 mV in cHet, n=40 cells from 17 mice, ***p<0.001, LMM; excluding the 2 outliers from cHet group -42.6±1.01 mV in control, n=33 cells from 15 mice vs -36.2±1.1 mV in cHet, n=38 cells from 17 mice, ***p<0.001, LMM.

Fig.7c: including the 2 outliers from cHet group with values of -16.5 and 20.6 mV: 43.4±1.6 mV in control, n=12 cells from 9 mice vs -33.9±1.8 mV in cHet, n=24 cells from 13 mice, **p=0.002, LMM; excluding the 2 outliers from cHet group -43.4±1.6 mV in control, n=12 cells from 9 mice vs -35.4±1.7 mV in cHet, n=22 cells from 13 mice, *p=0.037, LMM.

(5) The authors claim that "cHet SST+ cells showed no significant changes in active and passive membrane properties (Figure 8d,e); however, their evoked firing properties were affected with fewer AP generated in response to the same depolarizing current injection".

This sentence is intrinsically contradictory. Action potentials triggered by current injections are dependent on the integration of passive and active properties. If the curves of Figure 8f are different between genotypes, then some passive and/or active property MUST have changed. It is an unescapable conclusion. The general _blanket_ statement of the authors that there are no significant changes in active and passive properties is in direct contradiction with the current/#AP plot.

We agreed with the reviewer and rephrased the abstract, results and discussion according to better represent the data. As discussed in the previous revision, it's possible that other intrinsic factors, not assessed in this study, may have contributed to the effect shown in the current/#AP plot.

(6) The phase plots of Figs 5c, 7c, and 7h suggest that the frequency of acquisition/filtering of current-clamp signals was not appropriate for fast waveforms such as spikes. The first two papers indicated by the authors in their rebuttal (Golomb et al., 2007; Stevens et al., 2021) did not perform a phase plot analysis (like those included in the manuscript). The last work quoted in the rebuttal (Zhang et al., 2023) did perform phase plot analysis, but data were digitized at a frequency of 20KHz (not 10KHz as incorrectly indicated by the authors) and filtered at 10 kHz (not 2-3 kHz as by the authors in the manuscript). To me, this remains a concern.

We agree with the reviewer that higher sampling rate would allow to more accurately assess different AP parameters, such as AP peak, half-width, rise time, etc. The papers were cited in context of determining AP threshold, not performing phase plot analysis. We apologize for the confusion and error. Finally, we removed the phase plots since they did not add relevant information.

(7) The general logical flow of the manuscript could be improved. For example, Fig 4 seems to indicate no morphological differences in the dendritic trees of control vs mutant PV cells, but this conclusion is then rejected by Fig 6. Maybe Fig 4 is not necessary. Regarding Fig 6, did the authors check the integrity of the entire dendritic structure of the cells analyzed (i.e. no dendrites were cut in the slice)? This is critical as the dendritic geometry may affect the firing properties of neurons (Mainen and Sejnowski, Nature, 1996).

As suggested by the reviewer, we removed Fig.4. All the reconstructions used for dendritic analysis contained intact cells with no evidently cut dendrites.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: SYNGAP1 haploinsufficiency-related intellectual disability (SYNGAP1-ID) is characterized by moderate to severe ID, generalized epilepsy, autism spectrum disorder, sensory processing dysfunction and other behavioral abnormalities. While numerous studies have highlighted a role of Syngap1 in cortical excitatory neurons development; recent studies suggest that Syngap1 plays a role in GABAergic inhibitory neuron development as well. However, the molecular pathways by which Syngap1 acts on GABAergic neurons, and whether they are similar or different from the mechanisms underlying its effects in excitatory neurons, is unknown. Here, we examined whether, and how, embryonic-onset Syngap1 haploinsufficiency restricted to GABAergic interneurons derived from the medial ganglionic eminence (MGE) impacts their synaptic and intrinsic properties in adult primary auditory cortex (A1). We found that Syngap1 haploinsufficiency affects the intrinsic properties, overall leading to increased firing threshold, and decreased excitatory synaptic drive of Parvalbumin (PV)+ neurons in adult Layer IV A1, whilst Somatostatin (SST)+ interneurons were mostly resistant to Syngap1 haploinsufficiency. Further, the AMPA component of thalamocortical evoked-EPSC was decreased in PV+ cells from mutant mice. Finally, we found that the selective blocking of voltage-gated D-type K+ currents was sufficient to rescue PV+ mutant cell-intrinsic properties to wild-type levels. Together, these data suggest that Syngap1 plays a specific role in the maturation of PV+ cell intrinsic properties and synaptic drive, and its haploinsufficiency may lead to reduced PV cell recruitment in the adult A1, which could in turn contribute to the auditory processing alterations found in SYNGAP1-ID preclinical models and patients.

Revised version: SYNGAP1 haploinsufficiency-related intellectual disability (SYNGAP1-ID) is characterized by moderate to severe ID, generalized epilepsy, autism spectrum disorder, sensory processing dysfunction and other behavioral abnormalities. While numerous studies have highlighted a role of Syngap1 in cortical excitatory neurons development, recent studies suggest that Syngap1 plays a role in GABAergic inhibitory neuron development as well. However, the molecular pathways by which Syngap1 acts on GABAergic neurons, and whether they are similar or different from the mechanisms underlying its effects in excitatory neurons, is unknown. Here, we examined whether, and how, embryonic-onset Syngap1 haploinsufficiency restricted to GABAergic interneurons derived from the medial ganglionic eminence (MGE) impacts their synaptic and intrinsic properties in adult primary auditory cortex (A1). We found that Syngap1 haploinsufficiency significantly affected the intrinsic properties, overall leading to increased firing threshold, and decreased excitatory synaptic drive of Parvalbumin (PV)+ neurons in adult Layer IV A1. Further, the AMPA component of thalamocortical evoked-EPSC was decreased in PV+ cells from mutant mice. Mutant somatostatin (SST)+ interneurons exhibited decreased spontaneous excitatory input and impaired evoked firing without alterations in firing threshold. Finally, we found that the selective blocking of voltage-gated D-type K+ currents was sufficient to rescue PV+ mutant cell-intrinsic properties to wild-type levels. Together, these data suggest that Syngap1 plays a specific role in the maturation of PV+ cell intrinsic properties and synaptic drive, and its haploinsufficiency may lead to reduced PV cell recruitment in the adult A1, which could in turn contribute to the auditory processing alterations found in SYNGAP1-ID preclinical models and patients.

Introduction

Earlier version: Haploinsufficiency of SYNGAP1, an increasingly well-recognised candidate gene in neurodevelopmental disorders, causes intellectual disability (SYNGAP1-ID), epilepsy, autism spectrum disorder (ASD), sensory processing dysfunctions, including in the auditory domain, and other behavioral abnormalities (Hamdan et al., 2009; Berryer et al., 2013; Carreño-Muñoz et al., 2022). SYNGAP1 is an abundant synaptic GTPase-activating protein (GAP) critical for synaptic plasticity, learning, memory, and cognition (Kim et al., 1998; Gamache et al., 2020). Within the brain, Syngap1 expression is highest, but not limited to, forebrain structures, including the cortex and hippocampus (Kim et al., 1998; Porter et al., 2005), where it peaks during periods of robust synaptogenesis (Porter et al., 2005; McMahon et al., 2012; Gou et al., 2020; Jadhav et al, 2024). The function of Syngap1 has been mostly studied in excitatory neurons; in particular, Syngap1 haploinsufficiency has been shown to increase AMPA receptor density and cause premature maturation of excitatory synapses formed by hippocampal and somatosensory layer 5 pyramidal cells in rodents (Clement et al., 2012, 2013; Ozkan et al., 2014; Aceti et al, 2015). Consistently, experiments involving xenotransplantation of human cortical neurons into the mouse brain showed that SYNGAP1-deficient human neurons display accelerated synaptic formation and maturation alongside disrupted synaptic plasticity (Vermaercke et al., 2024). Syngap1 can also play a non-synaptic function linked to the control of cortical neurogenesis of projecting neurons (Birtele et al., 2023). In addition to excitatory neurons, Syngap1 mRNA and protein expression is also detected in inhibitory neurons (Zhang et al., 1999; Moon et al., 2008; Berryer et al., 2016; Su et al., 2019; Velmeshev et al., 2019; Zhao and Kwon, 2023; Jadhav et al, 2024). Further, recent studies have implicated Syngap1 in GABAergic cell migration and inhibitory synapse maturation (Berryer et al., 2016; Su et al., 2019; Sullivan et al., 2020; Khlaifia et al., 2023); nevertheless, whether the underlying cellular and molecular mechanisms are similar to those observed in excitatory cells is unknown.

Revised version: SYNGAP1 is a key synaptic GTPase-activating protein (GAP) essential for synaptic plasticity, learning, memory, and cognition (Kim et al., 1998; Gamache et al., 2020). Its expression is abundant within forebrain structures, including the cortex and hippocampus (Kim et al., 1998; Porter et al., 2005), where it peaks during critical periods of synaptogenesis (Porter et al., 2005; McMahon et al., 2012; Gou et al., 2020; Jadhav et al, 2024). SYNGAP1 is increasingly recognised as a candidate gene in neurodevelopmental disorders, with haploinsufficiency leading to intellectual disability (SYNGAP1-ID), epilepsy, autism spectrum disorder (ASD), sensory processing deficits, including in the auditory domain, and other behavioral abnormalities (Hamdan et al., 2009; Berryer et al., 2013; Carreño-Muñoz et al., 2022). The role of Syngap1 has been studied primarily in excitatory neurons. Specifically, Syngap1 haploinsufficiency has been shown to increase AMPA receptor density and accelerate the maturation of excitatory synapses in hippocampal and somatosensory layer 5 pyramidal cells in rodents (Clement et al., 2012, 2013; Ozkan et al., 2014; Aceti et al, 2015). Similarly, xenotransplantation experiments have shown that SYNGAP1-deficient human cortical neurons transplanted into mouse brains exhibit faster synapse formation and maturation, along with disrupted synaptic plasticity (Vermaercke et al., 2024). In addition to its synaptic roles, Syngap1 is also implicated in the regulation of cortical neurogenesis of projecting neurons (Birtele et al., 2023).

Added in the revised version: Although Syngap1 research has predominantly focused on excitatory neurons, its mRNA and protein are also expressed in inhibitory neurons (Zhang et al., 1999; Moon et al., 2008; Berryer et al., 2016; Su et al., 2019; Velmeshev et al., 2019; Zhao and Kwon, 2023; Jadhav et al, 2024). Emerging evidence suggests a role for Syngap1 in the migration of GABAergic cells and the maturation of inhibitory synapses (Berryer et al., 2016; Su et al., 2019; Sullivan et al., 2020; Khlaifia et al., 2023); nevertheless, whether the underlying cellular and molecular mechanisms are similar to those observed in excitatory cells is unknown.

Earlier version: Based on anatomical, physiological and specific marker expression, we can identify two major sub-classes of cortical inhibitory neurons, parvalbumin-(PV+) and somatostatin- (SST+) expressing interneurons, which provide perisomatic and distal dendritic inhibition to pyramidal cells, respectively (Levy and Reyes, 2012; Yavorska and Wehr, 2016). Due to the specificity in their synapse location and distinct functional properties, SST+ and PV+ neurons differentially shape excitatory neuronal responses. By providing fast inhibition onto postsynaptic pyramidal cell somata and proximal dendrites, PV+ cells exert fine control on their output (Moore and Wehr, 2013; Tremblay et al., 2016), while SST+ cells targeting apical dendrites of postsynaptic pyramidal cells exert specific control over dendritic synaptic integration (Kawaguchi and Kubota, 1997; Chiu et al., 2013). We recently showed that prenatal-onset Syngap1 haploinsufficiency restricted to Nkx2.1- expressing GABAergic interneuron precursors, which include PV+ and SST+ interneurons, leads to the development of alterations in auditory cortex activity, which resemble those observed in global Syngap1 haploinsufficient mouse models and SYNGAP1-ID patients (Carreño-Muñoz et al., 2022), suggesting that interneuron dysfunction may contributes to these specific phenotypes. However, how prenatal-onset Syngap1 haploinsufficiency in GABAergic interneurons alters their physiology in adult cortex is unknown.

Revised version: Cortical inhibitory neurons can be broadly classified into two major subtypes based on their anatomy, physiology and expression of specific markers: parvalbumin- (PV+) and somatostatin- (SST+) expressing interneurons (Rudy et al., 2011). These subtypes provide perisomatic and distal dendritic inhibition to pyramidal cells, respectively (Levy and Reyes, 2012; Yavorska and Wehr, 2016). Due to the specificity in their synapse location and distinct functional properties, SST+ and PV+ neurons differentially shape excitatory neuronal responses. By providing fast inhibition onto postsynaptic pyramidal cell somata and proximal dendrites, PV+ cells exert fine control on their output (Moore and Wehr, 2013; Tremblay et al., 2016), while SST+ cells targeting apical dendrites of postsynaptic pyramidal cells exert specific control over dendritic synaptic integration (Kawaguchi and Kubota, 1997; Chiu et al., 2013). We recently showed that prenatal-onset Syngap1 haploinsufficiency restricted to Nkx2.1-expressing GABAergic interneuron precursors, which include PV+ and SST+ interneurons, leads to the development of alterations in auditory cortex activity (Jadhav et al., 2024), which resemble those observed in global Syngap1 haploinsufficient mouse models and SYNGAP1-ID patients (Carreño-Muñoz et al., 2022), suggesting that interneuron dysfunction may contributes to these specific phenotypes. However, how prenatal-onset Syngap1 haploinsufficiency in GABAergic interneurons alters their physiology in adult cortex is unknown.

Earlier version: To address this question, we generated conditional transgenic mice wherein Syngap1 haploinsufficiency was restricted to MGE-derived interneurons. We then assessed the synaptic and intrinsic properties of PV+ fast spiking (FS) and SST+ regular spiking cells in layer IV of the adult primary auditory cortex (A1). We performed whole-cell voltage clamp recording in combination with electrical stimulation of thalamic fibers and current clamp recording, to study synaptic and intrinsic properties, respectively. We found that both mutant PV+ and SST+ cells show decreased excitatory synaptic drive; however, PV+, but not SST+, interneurons showed significantly increased threshold for action potential (AP) generation, which points towards a reduced recruitment of cortical PV cells in the mutant mouse. Further, we rescued PV+ cell hypofunction ex vivo using alpha-dendrotoxin (α-DTX), a drug specifically targeting the Kv1 group of voltage-gated D-type K+ currents. Altogether our results suggest that Syngap1 can affect neuronal physiological properties by modulating distinct molecular factors in a neuron-specific fashion.

Revised version: To address this question, we generated conditional transgenic mice wherein Syngap1 haploinsufficiency was restricted to MGE-derived interneurons. We then assessed the synaptic and intrinsic properties of PV+ fast spiking (FS) and SST+ regular spiking cells in layer IV of the adult primary auditory cortex (A1). We found that both mutant PV+ and SST+ cells show decreased excitatory synaptic drive. Notably, PV+, but not SST+, interneurons showed a significantly increased threshold for action potential (AP) generation, pointing towards a reduced recruitment of cortical PV cells in the mutant mouse. Further, we were able to partially restore PV+ cell function ex vivo using alpha-dendrotoxin (α-DTX), a selective blocker of Kv1 family voltage-gated D-type K+ currents. Overall, these results suggest that Syngap1 can affect neuronal physiological properties by modulating distinct molecular mechanisms in a cell type-specific manner.

Results

Syngap1 haploinsufficiency in MGE-derived interneurons is associated with decreased synaptic excitation in PV+ cells

Earlier version: Since Syngap1 has been shown to regulate AMPAR-mediated synaptic transmission in hippocampal and cortical excitatory neurons and hippocampal inhibitory neurons (Ozkan et al., 2014; Arora et al., 2022; Khlaifia et al., 2023), we first explored whether embryonic-onset Syngap1 haploinsufficiency in MGE-derived cortical interneurons affects their glutamatergic synaptic inputs in adult A1. We performed targeted voltage-clamp recordings of spontaneous (sEPSCs) and miniature excitatory postsynaptic currents (mEPSCs) in layer IV (LIV) EGFP-expressing interneurons from auditory thalamocortical slices of 9-13 weeks-old Tg(Nkx2.1-Cre):RCEf/f:Syngap1+/+(control) and Tg(Nkx2.1-Cre):RCEf/f:Syngap1f/+(cHet) littermates (Figure 1a–j, Figure supplement 1a-e, Table 1, Table S1). Nkx2.1-expressing MGE precursors generate most of PV+ and SST+ cortical interneurons (Xu et al., 2008), while the RCE allele drive Cre-dependent EGFP expression. Recorded MGE-derived interneurons were filled with biocytin, and their identity was confirmed by immunolabeling for neurochemical markers (PV or SST) and analysis of anatomical properties, such as axonal arborisation location and presence or absence of dendritic spines (Kawaguchi and Kubota, 2006; Rock et al., 2018; Bertero et al., 2020). We found that sEPSC amplitude, but not inter-event interval, was decreased in both LIV PV+ and SST+ neurons from control and cHet littermates (Figure 1b, Figure supplement 1b). sEPSC rise and decay time constants for both PV+ and SST+ interneurons was not significantly different between the two genotypes (Figure 1d, Figure supplement 1d), suggesting that postsynaptic AMPA receptor subunit composition was not affected by Syngap1 haploinsufficiency. Finally, charge transfer (area under sEPSC) and the product of mean PSC charge transfer and event frequency were both significantly decreased in SST+, but not PV+ cells in cHet mice compared to controls (Figure 1e, Figure supplement 1e). To discern whether Syngap1 haploinsufficiency had a pre or postsynaptic action on the glutamatergic drive received by PV+ cells, we recorded mEPSC (Figure 1f-j). We found no difference in mEPSC amplitude between the two genotypes (Fig. 1g), indicating that the observed difference in sEPSC amplitude (Figure 1b) could arise from decreased network excitability. Conversely, cHet mice showed a shift toward an increase in the inter-mEPSC time interval (Fig. 1g). Consistent with this observation, we observed a significant decrease in the density of putative glutamatergic synapses onto PV+ cell somata, identified by the colocalization of the vesicular glutamate 1 (vGlut1, presynaptic marker) and PSD95 (postsynaptic marker), in cHet compared to control mice (Figure 1, Figure supplement 2a, b). To test whether Syngap1 insufficiency affects excitatory inputs by reducing network activity, we recorded sEPSCs followed by mEPSCs from control and cHet PV+ neurons (Figure supplement 3a-d). In both genotypes, we found no significative difference in either amplitude or inter-event intervals between sEPSC and mEPSC, suggesting that in acute slices from adult A1, most sEPSCs may actually be AP-independent. While perhaps surprisingly at first glance, this result can be explained by recent published work suggesting that AP-dependent (sEPSC) and -independent (mEPSC) release may not necessarily engage the same pool of vesicles or target the same postsynaptic sites (Sara et al., 2005; Sara et al., 2011; Ramirez and Kavalali, 2011; Kavalali, 2015). Nevertheless, these data suggest that cHet PV+ cells receive reduced glutamatergic drive compared to control PV+ cells.

Revised version: Since Syngap1 has been shown to regulate AMPAR-mediated synaptic transmission in hippocampal and cortical excitatory neurons and hippocampal inhibitory neurons (Ozkan et al., 2014; Arora et al., 2022; Khlaifia et al., 2023), we first explored whether embryonic-onset Syngap1 haploinsufficiency in MGE-derived cortical interneurons affects their glutamatergic synaptic inputs in adult A1. We performed targeted voltage-clamp recordings of spontaneous (sEPSCs) and miniature excitatory postsynaptic currents (mEPSCs) in layer IV (LIV) EGFP-expressing interneurons from auditory thalamocortical slices of 9-13 weeks-old Tg(Nkx2.1- Cre):RCEf/f:Syngap1+/+(control) and Tg(Nkx2.1-Cre):RCEf/f:Syngap1f/+(cHet) littermates (Figure 1, Table 1). Nkx2.1-expressing MGE precursors generate most of PV+ and SST+ cortical interneurons (Xu et al., 2008), while the RCE allele drive Cre-dependent EGFP expression. Recorded MGE-derived interneurons were filled with biocytin, and their identity was confirmed by immunolabeling for neurochemical markers (PV or SST) and analysis of anatomical properties, including location of the axonal arborisation across different cortical layers and presence or absence of dendritic spines (Figure 1a, Suppl. Figure 1) (Kawaguchi and Kubota, 2006; Rock et al., 2018; Bertero et al., 2020). All results were analysed by linear mixed model (LMM), modelling animal as a random effect and genotype as the fixed effect. This method was chosen for statistical analysis because it accounts for both animals as independent replicates and cell recorded in each mouse as repeated/correlated measures, thus providing the most accurate approach for assessing the data (Aarts et al., 2014; Yu et al 2022).

Earlier version: Syngap1 haploinsufficiency in Nkx2.1+ interneurons is associated with reduced sEPSC amplitude and mEPSC frequency in LIV PV+ cells.

Revised version: Syngap1 haploinsufficiency in Nkx2.1+ interneurons is associated with reduced sEPSC amplitude and mEPSC frequency in LIV BCs.

Earlier version: (a) Representative traces of sEPSCs recorded in PV+ cells from control Tg(Nkx2.1-Cre):RCEf/f:Syngap1+/+(blue, n=28 cells, 11 mice) and cHet Tg(Nkx2.1-Cre):RCEf/f:Syngapf/+(red, n=28 cells, 11 mice) mice. (b), Cumulative probability plots show a significant decrease in the amplitude (LMM, p=0.048) and no change in the inter-sEPSC interval (LMM, p=0.186). Insets illustrate significant differences in the sEPSC amplitude for inter-cell mean comparison (LMM, *p=0.044) and no difference for inter-sEPSC interval (LMM, p=0.676). (c) Representative examples of individual sEPSC events (100 pale sweeps) and average traces (bold trace) detected in PV cells of control and cHet mice. (d) Superimposed scaled traces (left top) and summary bar graphs for a group of cells (bottom) show no significant differences in the sEPSCs kinetics (LMM, p=0.350 for rise time, and p=0.429 for decay time). (e) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM, p=0.064, left) and for the charge transfer when frequency of events is considered (LMM, p=0.220). (f) Representative traces of mEPSCs recorded in PV+ cells from control (blue, n=18 cells, 7 mice) and cHet (red, n=8 cells, 4 mice) mice. (g) Cumulative probability plots show no change in the amplitude of mEPSC (LMM, p=0.135) and an increase in the inter-mEPSC interval (LMM, *p=0.027). Insets illustrate summary data showing no significant differences in the amplitude (LMM, p=0.185) and the inter-mEPSCs interval for inter-cell mean comparison (LMM, p=0.354). (h) Representative examples of individual mEPSC events (100 pale sweeps) and average traces (bold trace) detected in PV+ cells of control and cHet mice. (i) Summary bar graphs for a group of cells show no significant differences in the quantal content (LMM, p=0.222) and mEPSCs kinetics (LMM, p=0.556 for rise time, and p=0.591 for decay time). (j) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM p=0.069, left) and for the charge transfer when frequency of events is considered (LMM p=0.137).

Revised version: (a) Anatomical reconstructions of PV+ cells filled with biocytin in control (left) and cHet mice (right) during whole-cell patch-clamp recordings and post hoc immunohistochemical validation of BC interneurons confirming the positivity for PV. (b) Representative traces of sEPSCs recorded in BCs cells from control Tg(Nkx2.1-Cre):RCEf/f:Syngap1+/+ (blue, n=14 cells, 7 mice) and cHet Tg(Nkx2.1-Cre):RCEf/f:Syngapf/+(red, n=11 cells, 6 mice) mice. (c) Cumulative probability plots show a significant decrease in the amplitude of sEPSC in cHet mice compared to control mice (hash sign denote the significance for LMM related to the cumulative distributions ,# p=0.029) and no change in the inter-sEPSC interval (LMM, p=0.345). Insets illustrate significant differences in the sEPSC amplitude for inter-cell mean comparison (LMM, *p=0.014) and no difference for inter-sEPSC interval (LMM, p=0.230). (d) Representative examples of individual sEPSC events (100 pale sweeps) and average traces (bold trace) detected in BCs of control and cHet mice and superimposed scaled traces (right top), (e) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM, p=0.090, left) and for the charge transfer when frequency of events is considered (LMM, p=0.140). (f) Representative traces of mEPSCs recorded from the same neurons shown in b-e. (g) Cumulative probability plots show no change in the amplitude of mEPSC (LMM, p=0.151) and a significant increase in the inter-mEPSC interval in cHet mice compared to control mice (LMM, #p=0.045). Insets illustrate summary data showing no significant differences in the amplitude (LMM, p=0.155) and a trend towards longer inter-mEPSCs intervals for inter-cell mean comparison in cHet compared to control mice (LMM, p=0.056). (h) Representative examples of individual mEPSC events (100 pale sweeps) and average traces (bold trace) detected in BCs of control and cHet mice. (i) Summary bar graphs for a group of cells show no significant differences in the quantal content (LMM, p=0.189) and mEPSCs kinetics (LMM, p=0.269 for rise time, and p=0.193 for decay time). (j) Summary bar graphs showing a significant decrease in cHet mice for inter-cell mean charge transfer (LMM, *p=0.012, left) and for the charge transfer when frequency of events is considered (LMM, **p=0.002). * and # indicates p value <0.05 for bar graphs and cumulative distribution, respectively.

Added in the revised version: We found that sEPSC amplitude was decreased in LIV PV+ neurons from cHet mice compared to those recorded in control littermates (Figure 1c right). sEPSC rise and decay time constants were not significantly different between the two genotypes (Table 1), suggesting that postsynaptic AMPA receptor subunit composition was not affected by Syngap1 haploinsufficiency. To discern whether Syngap1 haploinsufficiency had a pre or postsynaptic effect on the glutamatergic drive received by PV+ cells, we analysed mEPSC recorded from the same neurons shown in Figure 1b-e (Figure 1f- j). We confirmed that TTX blocked APs in PV+ cells and that recordings were stable as indicated by lack of changes in series resistance during the recording period in our experimental setup (Suppl. Figure 2f-i). We found no difference in mEPSC amplitude between the two genotypes (Fig. 1g, right), indicating that the observed difference in sEPSC amplitude (Figure 1c, right) could be due to impaired AP-dependent release in cHet mice and the presence of large-amplitude sEPSCs that are preferentially affected by TTX in control mice (Suppl. Figure 2b-e). Conversely, cHet mice showed longer inter-mEPSC time interval (Figure 1g, left, for inter-cell mean comparison LMM, p=0.056, for cumulative distributions LMM, p=0.045), and significantly lower charge transfer and ΔQ*f (Figure 1j) compared to controls littermates, suggesting a decrease of glutamatergic presynaptic release sites onto PV+ cells. Recent studies, based on serial block-face scanning electron microscopy, suggest that cortical PV+ interneurons receive more robust excitatory inputs to their perisomatic region than pyramidal neurons (Hwang et al., 2021; Elabbady et al., 2025). Quantification of the density of putative glutamatergic synapses onto PV+ cell somata, identified by the colocalization of the vesicular glutamate 1 (vGlut1, presynaptic marker) and PSD95 (postsynaptic marker), revealed a significant decrease in cHet compared to control mice (Suppl. Figure 3a,b). Whether the density of dendritic targeting excitatory inputs is affected as well remains an open question.

Earlier version: In A1, layer IV PV+ cells receive stronger subcortical thalamocortical inputs compared to excitatory cells and other subpopulations of GABAergic interneurons (Ji et al., 2016; Rock et al., 2018). We thus recorded evoked AMPA (eAMPA)- and NMDA (eNMDA)-mediated currents in PV+ cells by bulk electrical stimulation of the thalamic radiation (Figure 2a-f, Table 2), to determine whether thalamocortical synapses were affected by conditional Syngap1 haploinsufficiency. eAMPA amplitude, area under the curve (AUC) charge transfer (average of all responses, successes + failures, Figure 2b,c) and potency (average of all successes only, Figure 2e) were decreased in cHet mice as compared to control littermates. In addition, we found a substantial increase in onset latencies of eAMPA currents (Figure 2d), suggesting a potential deficit in the thalamocortical recruitment of PV+ cells. Next, we assessed eNMDA currents in PV+ cells in presence of GABAAR, GABAB and AMPA inhibitors (1 μM Gabazine, 2 μM CGP, and 10 μm NBQX, respectively) (Figure 2c,e). We found that eNMDA currents as well as the percentage of PV+ cells showing these responses were similar in cHet and control littermates (Figure 2c,e), therefore leading to increased NMDA/AMPA in cHet mice (Figure 2e). Interestingly, the kinetics of eAMPA and eNMDA currents were similar in both genotypes, indicating no change in their subunit composition (Figure 2f, Table 2). These results suggest that embryonic-onset Syngap1 haploinsufficiency in MGE-derived interneurons specifically impairs AMPA-mediated thalamocortical recruitment of PV+ cells. Of note, the density of putative thalamocortical glutamatergic synapses onto PV+ cell somata, identified by the colocalization of the vesicular glutamate 2 (vGlut2, thalamocortical presynaptic marker) and PSD95 (postsynaptic marker) was not significantly different in cHet as compared to littermate controls (Figure 2, Figure Supplement 2c,d), suggesting that presynaptic release from excitatory thalamocortical fibers or/and AMPARs expression at thalamocortical synapses on PV+ cells are likely decreased in cHet mice.

Revised version: In A1, layer IV PV+ cells receive stronger subcortical thalamocortical inputs compared to excitatory cells and other subpopulations of GABAergic interneurons (Ji et al., 2016; Rock et al., 2018). We thus recorded evoked AMPA (eAMPA)- and NMDA (eNMDA)-mediated currents in PV+ cells by bulk electrical stimulation of the thalamic radiation (Figure 2a-f, Table 2), to determine whether thalamocortical synapses were affected by conditional Syngap1 haploinsufficiency. eAMPA amplitude, area under the curve (AUC) charge transfer (average of all responses, successes + failures, Figure 2b right,c) and potency (average of all successes only, Figure 2e left) were decreased in cHet mice as compared to control littermates. In addition, we found a substantial increase in onset latencies of eAMPA currents (Figure 2d), suggesting a potential deficit in the thalamocortical recruitment of PV+ cells. Next, we assessed eNMDA currents in PV+ cells in presence of GABAAR, GABAB and AMPA inhibitors (1 μM Gabazine, 2 μM CGP, and 10 μm NBQX, respectively) (Figure 2c,e). We found that eNMDA currents as well as the fraction of PV+ cells showing these responses were similar in cHet and control littermates (Figure 2c right, 2e left), therefore leading to increased NMDA/AMPA in cHet mice (Figure 2e right). Interestingly, the kinetics of eAMPA and eNMDA currents were similar in both genotypes, indicating no change in their subunit composition (Figure 2f, Table 2). These results suggest that embryonic-onset Syngap1 haploinsufficiency in MGE-derived interneurons specifically impairs AMPA-mediated thalamocortical recruitment of PV+ cells. Of note, the density of putative thalamocortical glutamatergic synapses onto PV+ cell somata, identified by the colocalization of the vesicular glutamate 2 (vGlut2, thalamocortical presynaptic marker) and PSD95 (postsynaptic marker) was not significantly different in cHet as compared to littermate controls (Suppl. Figure 3c,d), suggesting that presynaptic release from excitatory thalamocortical fibers or/and AMPARs expression at thalamocortical synapses on PV+ cells are likely decreased in cHet mice.

Thalamocortical eAMPA transmission is decreased in LIV PV+ cells from Tg(Nkx2.1-Cre):RCEf/f:Syngap1f/+mice.

Earlier version: (a) Representative examples of individual eAMPA (negative deflections) and eNMDA (positive deflections) (5-10 pale sweep) and average traces (bold trace) recorded in PV cells from control (blue, n=16 cells, 7 mice) and cHet mice (red, n=14 cells, 7 mice). (b) Summary plots showing no change in the failure rate of eAMPA (left, LMM, p=0.550) and a significant increase in the minimal (including failures and successes) eAMPA amplitude (LMM, *p=0.031) and (c, left) charge transfer (LMM, *p=0.033) in cHet mice. (c, right) Summary bar graph illustrating the percentage of PV+ cells containing eNMDA in the thalamocortical evoked EPSC. (d) Synaptic latency histograms (bottom) of thalamocortical eEPSC from control and cHet mice, and summary bar graph (top) illustrating an increase in the onset latencies of eEPSC in cHet mice (LMM, *p=0.023). For both histograms, bins are 0.1 msec wide. (e) Summary plots showing a significant increase in the potency of eAMPA (successes only, LMM, **p=0.003, left) in cHet mice with no change in eNMDA (LMM, p=0.969), A significant increase is present in the NMDA/AMPA ratio (i.e. ratio of the peak for eNMDA and eAMPA, LMM, **p=0.001, right) in cHet mice. (f) Summary plots showing no significant differences in the eAMPA (LMM, p=0.177 for rise time, and p=0.608 for decay time) and eNMDA kinetics (LMM, p=0.228 for rise time, and p=0.221 for decay time). (g) Representative examples of individual LIV eEPSC (10 pale sweeps) and average traces (bold trace) with an interval of 50 ms recorded in two BC cells from control (blue) and cHet mice (red). (h) Summary plot showing significantly increased PPRs recorded from LIV BC cHet (red circles, n=8 cells; 5 mice) compared to controls (blue circles, n=9 cells, 5 mice), when two EPSCs were evoked in layer IV BC with two electric pulses at 30 or 50ms intervals (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, **p=0.001).

Revised version: (a) Representative examples of individual eAMPA (negative deflections) and eNMDA (positive deflections) (5-10 pale sweep) and average traces (bold trace) recorded in PV cells from control (blue, n=16 cells, 7 mice) and cHet mice (red, n=14 cells, 7 mice). (b, left) Summary plots showing no change in the failure rate of eAMPA (left, LMM, p=0.550) and a significant decrease in the minimal (including failures and successes) eAMPA amplitude (b, right LMM, *p=0.031) and (c, left) charge transfer (LMM, *p=0.033) in cHet mice. (c, right) Summary bar graph illustrating the percentage of PV+ cells containing eNMDA in the thalamocortical evoked EPSC. (d) Synaptic latency histograms (bottom) of thalamocortical eEPSC from control and cHet mice, and summary bar graph (top) illustrating an increase in the onset latencies of eEPSC in cHet mice (LMM, *p=0.023). For both histograms, bins are 0.1 msec wide. (e) Summary plots showing a significant decrease in the potency of eAMPA (successes only, LMM, **p=0.003, left) in cHet mice with no change in eNMDA (LMM, p=0.969), A significant increase is present in the NMDA/AMPA ratio (i.e. ratio of the peak for eNMDA and eAMPA, LMM, **p=0.001, right) in cHet mice. (f) Summary plots showing no significant differences in the eAMPA (LMM, p=0.177 for rise time, and p=0.608 for decay time) and eNMDA kinetics (LMM, p=0.228 for rise time, and p=0.221 for decay time). (g) Representative examples of individual LIV evoked EPSC (10 pale sweeps) and average traces (bold trace) with an interval of 50 ms recorded in two BC cells from control (blue) and cHet mice (red). (h) Summary plot showing significantly increased PPRs recorded from LIV BC cHet (red circles, n=8 cells; 5 mice) compared to controls (blue circles, n=9 cells, 5 mice), when two EPSCs were evoked in layer IV BC with two electric pulses at 30 or 50ms intervals (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, **p=0.001). * indicates p value <0.05; ** indicate p value <0.005.

Earlier version: To further address this issue, we performed paired pulse ratio (PPR) experiments, as in this case we are comparing the ability of the thalamic fibers reaching PV+ cells to release the same pool of vesicles (Figure 2g,h). We found that, in contrast with Control mice, evoked excitatory inputs to layer IV PV+ cells showed paired-pulse facilitation in cHet mice (Figure 2g, h), suggesting that thalamocortical presynaptic sites likely have decreased release probability in mutant compared to control mice.

Revised version: To further explore this phenomenon, we performed paired pulse ratio (PPR) experiments, as PPR typically reflects changes in presynaptic release probability (Figure 2g,h). We found that, in contrast with Control mice, evoked excitatory inputs to layer IV PV+ cells showed paired-pulse facilitation in cHet mice (Figure 2g,h), suggesting that thalamocortical presynaptic sites likely have decreased release probability in mutant compared to control mice.

Earlier version: Since PV+ cell recruitment is regulated by the balance of its excitatory and inhibitory inputs, we next analysed spontaneous (sIPSCs) and miniature inhibitory postsynaptic currents (mEPSCs) recorded from layer IV PV+ cells in both genotypes. We observed reduced sIPSC amplitude, but no significant changes in frequency and kinetics, in cHet compared to control PV+ cells (Figure 3a-e, Table 3). However, mIPSC analysis revealed no genotype-dependent differences in any parameters (Figure 3a-j, Table 3), suggesting that decreased sIPSC amplitude in cHet PV+ cells was likely due to changes in presynaptic cell-intrinsic excitability and/or network activity.

Revised version: Since PV+ cell recruitment is regulated by the balance of its excitatory and inhibitory inputs, we next analysed spontaneous (sIPSCs) and miniature inhibitory postsynaptic currents (mIPSCs) recorded from layer IV PV+ cells in both genotypes. We observed reduced sIPSC amplitude in cHet compared to control PV+ cells (Figure 3a,b, Table 3); however, mIPSC analysis revealed no genotype-dependent differences in any parameters (Figure 3f-j, Table 3), suggesting that decreased sIPSC amplitude in cHet PV+ cells was likely due to changes in presynaptic cell-intrinsic excitability and/or network activity.

The amplitude of sIPSCs, but not mIPSCs, in LIV PV+ cells is reduced in Tg(Nkx2.1-Cre):RCEf/f:Syngap1f/+mice

Earlier version: (a) Representative traces of sIPSCs recorded in PV+ cells from control (blue, n=27 cells, 8 mice) and cHet (red, n=24 cells, 7 mice) mice. (b), Cumulative probability plots show a significant decrease in the amplitude (LMM, **p=0.003) and no change in the inter-sIPSC interval (LMM, p=0.069). Insets illustrate significant differences in the sIPSC amplitude for inter-cell mean comparison (LMM, **p=0.002) and no difference for inter-sIPSC interval (LMM, p=0.102). (c) Representative examples of individual sIPSC events (100 pale sweeps) and average traces (bold trace) detected in PV cells of control and cHet mice. (d) Summary bar graphs for a group of cells show no significant differences in the sIPSCs kinetics (LMM, p=0.113 for rise time, and p=0.602 for decay time). (e) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM, p=0.234, left) and for the charge transfer when frequency of events is considered (LMM, p=0.273). (f) Representative traces of mIPSCs recorded in PV+ cells from control (blue, n=25 cells, 8 mice) and cHet (red, n=24 cells, 7 mice) mice. (g) Cumulative probability plots show no change in the amplitude of mIPSC (LMM, p=0.118) and in the inter-mIPSC interval (LMM, p=411). Insets illustrate summary data showing no significant differences in the amplitude (LMM, p=0.195) and the inter-mIPSCs interval for inter-cell mean comparison (LMM, p=0.243). (h) Representative examples of individual mIPSC events (100 pale sweeps) and average traces (bold trace) detected in PV+ cells of control and cHet mice. (i) Summary bar graphs for a group of cells show no significant differences in the mIPSCs kinetics (LMM, p=0.103 for rise time, and p=0.597 for decay time). (j) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM, p=0.374, left) and for the charge transfer when frequency of events is considered (LMM, p=0.100).

Revised version: (a) Representative traces of sIPSCs recorded in PV+ cells from control (blue, n=25 cells, 8 mice) and cHet (red, n=24 cells, 7 mice) mice. (b) Cumulative probability plots show a significant decrease in the amplitude of sIPSC in cHet mice compared to control (LMM, ## p=0.003) and no change in the inter-sIPSC interval (LMM, p=0.106). Insets illustrate significant differences in the sIPSC amplitude for inter-cell mean comparison (LMM, *p=0.009) and no difference for inter-sIPSC interval (LMM, p=0.185). (c) Representative examples of individual sIPSC events (100 pale sweeps) and average traces (bold trace) detected in PV cells of control and cHet mice. (d) Summary bar graphs for a group of cells show no significant differences in the sIPSCs kinetics (LMM, p=0.113 for rise time, and p=0.602 for decay time). (e) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM, p=0.234, left) and for the charge transfer when frequency of events is considered (LMM, p=0.273). (f) Representative traces of mIPSCs recorded in PV+ cells from control (blue, n=25 cells, 8 mice) and cHet (red, n=24 cells, 7 mice) mice. (g) Cumulative probability plots show no change in the amplitude of mIPSC (LMM, p=0.118) and in the inter-mIPSC interval (LMM, p=0.411). Insets illustrate summary data showing no significant differences in the amplitude (LMM, p=0.195) and the inter-mIPSCs interval for inter- cell mean comparison (LMM, p=0.243). (h) Representative examples of individual mIPSC events (100 pale sweeps) and average traces (bold trace) detected in PV+ cells of control and cHet mice. (i) Summary bar graphs for a group of cells show no significant differences in the mIPSCs kinetics (LMM, p=0.103 for rise time, and p=0.597 for decay time). (j) Summary bar graphs showing no differences for inter-cell mean charge transfer (LMM, p=0.374, left) and for the charge transfer when frequency of events is considered (LMM, p=0.100). * indicates p value <0.05 for bar graphs. # # indicates p value <0.005 for cumulative distribution.

Removed in the revised version: Layer IV PV+ cells show normal dendritic morphology in Nkx2.1 Cre+/-RCEf/f Syngap1f/+ mice

Removed in the revised version: Reduced glutamatergic drive onto PV+ cells may be due to impaired dendritic development. Since Syngap1 haploinsufficiency has been shown to impact the dendritic arbor of glutamatergic neurons (Clement et al., 2012; Michaelson et al., 2018; Arora et al., 2022), we next asked whether embryonic-onset conditional Syngap1 haploinsufficiency in PV+ cells had similar effects. We reconstructed the full dendritic arbor of PV+ cells filled with biocytin during patch-clamp recordings followed by posthoc immunolabeling with anti-PV and anti-SST antibodies (Figure 4a-c, Table 4). We observed that PV+ cells had an ovoid somata and multipolar dendrites typical of the basket cell (BC) family (Figure 4a). We found no genotype-dependent statistically significant differences in soma perimeter, number of branching points, total dendritic length and total dendritic area (Figure 4b). Further, dendritic Sholl analysis revealed no differences in dendritic intersections at different distances from the soma (Figure 4c). Together, these data indicate that in contrast to observations in cortical glutamatergic neurons (Clement et al., 2012; Michaelson et al., 2018; Arora et al., 2022), embryonic-onset Syngap1 haploinsufficiency in PV+ cells did not alter their dendritic arborisation. Therefore, structural modifications cannot explain the observed reduction in their glutamatergic drive (Figure 1 a-j).

Removed in the revised version: Syngap1 haploinsufficiency in Nkx2.1+ interneurons does not appear to alter the dendritic arbour of PV+ cells.

Removed in the revised version: (a) Anatomical reconstructions of PV+ cells filled with biocytin in control (left, n=10 cells, 6 mice) and cHet (right, n=12 cells, 7 mice) during whole-cell patch-clamp recordings and post hoc immunohistochemical validation of BC interneurons confirming the positivity for PV. (b) Summary data showing no differences in somatic (LMM, p=0.636) and dendritic parameters between the two genotypes (LMM, p=0.456 for terminals #, p=0.887 for total dendritic length, p=0.100 for branching points # and p=0.460 for dendritic surface area). (c, left) Representative reconstruction of PV+ cells and superimposed concentric circles used for Sholl analysis, with a radius interval of 10 µm from the soma. (c, center) Sholl analysis of PV+ cells dendritic branch patterns revealed no differences in the number of dendritic intersections (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, p=0.937). (c, right) Summary bar graph showing no differences in the total number of dendritic intersections obtained from Sholl analysis in control and cHet mice (LMM, p=0.502).

cHet mice show decreased layer IV PV+ cell-intrinsic excitability

Earlier version: PV+ cell recruitment in cortical circuits is dependent on both their synaptic drive and intrinsic excitability. Thus, we sought to investigate how Syngap1 haploinsufficiency in MGE-derived interneurons impacts the intrinsic excitability and firing properties of PV+ cells, by performing whole-cell current-clamp recordings (Figure 5a-e, Table 5). In line with preserved neuronal morphology, we found no changes in passive membrane properties (Cm, Rin and τ) of PV+ cells recorded from cHet mice as compared to control littermates (Figure 5a). However, analysis of active membrane properties revealed a significant decrease in the excitability of mutant PV+ cells (Figure 5b,c). In particular, cHet PV+ cells showed reduced AP amplitude, and increased AP threshold and latency to first AP (Figure 5b,c), while overall AP kinetics, including AP rise and decay time, fAHP time and AP half-width were not affected (Figure 5b, Table 5). In line with the decrease in intrinsic excitability, the rheobase (the smallest current injection that triggers an AP) was increased in PV+ cells from cHet mice (Figure 5d). In addition, both cHet and control PV+ cells displayed typical, sustained high-frequency trains of brief APs with little spike frequency adaptation in response to incremental current injections (Fig. 5e, right). However, cHet PV+ interneurons fired significantly fewer APs in response to the same depolarizing current injection when compared to control mice (Figure 5e, left). These data show that embryonic-onset Syngap1 haploinsufficiency in PV+ cells impairs their basic intrinsic and firing properties.

Revised version: PV+ cell recruitment in cortical circuits is dependent on both their synaptic drive and intrinsic excitability. Thus, we sought to investigate how Syngap1 haploinsufficiency in MGE-derived interneurons impacts the intrinsic excitability and firing properties of PV+ cells, by performing whole-cell current-clamp recordings (Figure 4, Table 4). We found no changes in passive membrane properties (Cm, Rin and τ) of PV+ cells recorded from cHet mice as compared to control littermates (Figure 4a); conversely, analysis of active membrane properties revealed a significant decrease in the excitability of mutant PV+ cells (Figure 4b,c). In particular, cHet PV+ cells showed reduced AP amplitude, and increased AP threshold and latency to first AP (Figure 4b,c, Table 4). In line with the decrease in intrinsic excitability, the threshold current was increased in PV+ cells from cHet mice (Figure 4d). Both cHet and control PV+ cells displayed typical, sustained high- frequency trains of brief APs with little spike frequency adaptation in response to incremental current injections (Fig. 4e, right); however, cHet PV+ interneurons fired significantly fewer APs in response to the same depolarizing current injection when compared to control mice (Figure 4e, left; Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ****p<0.0001). These data show that embryonic-onset Syngap1 haploinsufficiency in PV+ cells impair their basic intrinsic and firing properties.

PV+ cells intrinsic excitability is decreased in Tg(Nkx2.1-Cre):RCEf/f:Syngap1f/+ mice.

Earlier version: (a) Summary data showing no changes in the passive membrane properties between control (blue, n=33 cells, 15 mice) and cHet mice (red, n=40 cells, 17 mice) (LMM, p=0.081 for Vm, p=0.188 for Rin, p=0.188 for Cm, p=0.199 for Tm) (b) Summary data showing no differences in AP half-width (LMM, p=0.111) but a significant decrease in AP amplitude (LMM, *p=0.032) and a significant increase in AP latency (LMM, *p=0.009) from PV+ cells recorded in cHet mice. (c, left) Summary bar graph shows a significant decrease in AP threshold from cHet mice (LMM, ***p<0.001) and phase plane (Vm versus dV/dt, c, right bottom) for the first AP generated confirmed this result with a more hyperpolarized value (the area squared) of AP threshold for generation of AP. (c, right top) Representative single APs evoked by rheobase currents from control and cHet mice. APs are aligned at 50% of the rising phase on X axis and peak on Y axis. Note the more hyperpolarized AP with consequent reduction in AP amplitude in PV+ cells from cHet mice. (d) Summary bar graph shows a significant increase in the rheobase current (LMM, **p=0.004). (e, left) Summary plot showing a reduction of averaging number of APs per current step (40 pA) amplitude recorded from LIV PV+ cHet (red circles, n=38 cells; 18 mice) compared to control (blue circles, n=30 cells, 15 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ****p<0.0001). (e, right) Representative voltage responses indicating the typical FS firing pattern of PV+ cells in control and cHet mice in response to depolarizing (+120 pA and +240 pA) current injections corresponding to rheobase and 2x rheobase current.

Revised version: (a) Summary data showing no changes in the passive membrane properties between control (blue, n=33 cells, 15 mice) and cHet mice (red, n=40 cells, 17 mice) (LMM, p=0.081 for Vm, p=0.188 for Rin, p=0.188 for Cm, p=0.199 for τ). (b) Summary data showing no differences in AP half-width (LMM, p=0.111) but a significant decrease in AP amplitude (LMM, *p=0.032) and a significant increase in AP latency (LMM, *p=0.009) from PV+ cells recorded in cHet mice. (c, left) Summary bar graph shows a significant increase in AP threshold from cHet mice (LMM, ***p<0.001) for the first AP generated. (c, right top) Representative single APs evoked by threshold currents from control and cHet mice. APs are aligned at 50% of the rising phase on X axis and peak on Y axis. Note the more hyperpolarized AP with consequent reduction in AP amplitude in PV+ cells from cHet mice. (d) Summary bar graph shows a significant increase in the threshold current (LMM, **p=0.004). (e, left) Summary plot showing a reduction of averaging number of APs per current step (40 pA) amplitude recorded from LIV PV+ cHet (red circles, n=38 cells; 18 mice) compared to control (blue circles, n=30 cells, 15 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ****p<0.0001). (e, right) Representative voltage responses indicating the typical FS firing pattern of PV+ cells in control and cHet mice in response to depolarizing (+120 pA and +240 pA) current injections corresponding to threshold current and 2x threshold current. * indicates p value <0.05; *** indicates p value <0.001; **** indicates value p<0.0001.

Earlier version: Embryonic-onset Syngap1 haploinsufficiency in MGE-derived interneurons differentially affects the dendritic arbor and intrinsic properties of two distinct PV+ cell subpopulations

Revised version: Embryonic-onset Syngap1 haploinsufficiency in MGE-derived interneurons differentially impacts dendritic arborization and intrinsic properties of distinct PV+ cell subpopulations

Earlier version: The majority of PV+ cells are classified as FS cells, due to their ability to sustain high-frequency discharges of APs (Figure 5e, right). However, clusters of atypical PV+ cells have been previously reported in several brain regions including subiculum (Nassar et al., 2015), striatum (Bengtsson Gonzales et al., 2020), hippocampus (Ekins et al., 2020) and somatosensory cortex (Helm et al., 2013). Atypical PV+ cells share many electrophysiological parameters with FS cells; however, they have a slower AP half-width and possess a lower maximal AP firing frequency (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020;). Based on these 2 criteria, in our data we found a moderate negative correlation between Fmaxinitial and AP half-width in both genotypes (Figure 6a, left), indicating the possibility that PV+ cell with broad AP duration could have a lower Fmaxinitial, a feature previously observed in atypical PV+ cells from other cortical areas (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020). To further investigate whether A1 PV+ cells from control mice could be functionally segregated into distinct clusters, we performed hierarchical clustering of AP half-width and Fmaxinitial values based on Euclidean distance (Figure 6a, right). Hierarchical clustering builds a map (dendrogram) quantifying the similarity between samples (PV+ cell) and clusters of samples. Based on this analysis, we identified 2 clusters of PV+ cells, with short AP half-widths associated with higher values of Fmaxinitial. Despite hierarchical clustering defining 2 subgroups of PV+ cells, there were still few PV+ cells with longer AP half-widths showing Fmaxinitial values typical of PV+ cells with shorter AP half-widths (Figure 6a, right), indicating that these 2 parameters alone may not enough to segregate our database of PV+ cells in subtypes. We thus decided to perform PCA analysis using additional key intrinsic physiological features such as passive (Vm, Rin, Cm) and active (rheobase, AP half-width, AP amplitude, first AP latency, AP threshold, fAHP amplitude, amplitude AR, frequency AR, Fmaxinitial and Fss) membrane properties (Figure 6b, left). We then chose the intersection point of the 2 AP half-width distributions in control and cHet mice as a cutoff to define 2 different subpopulations of PV+ cells, which we termed Basket Cell (BC)-short (AP half-width <0.78ms) and BC-broad (AP half-width ≥0.78ms) (Figure 6c,d). In control mice, these 2 PV+ cell subtypes showed major differences in Rin, Cm, Fmaxinitial and Fss (Table S2, comparison between BC-short and BC-broad in control mice). Using PCA analysis we also noticed that, while in control mice we could clearly distinguish 2 PV+ cell subgroups, the differences were more ambiguous in cHet mice wherein some BC-short felt in the subdivision of BC-broad (Figure 6d). In addition, we observed that the percentage of BC-short was increased in cHet compared to control mice (61% vs 41% of total PV+ cells, respectively), suggesting that Syngap1 haploinsufficiency affects specific subgroups of PV+ cells (Figure 6d).

Revised version: The majority of PV+ cells are classified as FS cells, due to their ability to sustain high-frequency discharges of APs (Figure 4e, right). However, clusters of atypical PV+ cells have been previously reported in several brain regions including subiculum (Nassar et al., 2015), striatum (Bengtsson Gonzales et al., 2020), hippocampus (Ekins et al., 2020) and somatosensory cortex (Helm et al., 2013). While atypical PV+ cells share many electrophysiological parameters with FS cells, they have a slower AP half-width and possess a lower maximal AP firing frequency (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020;). Using these 2 criteria, we found a moderate negative correlation between Fmaxinitial and AP half-width in both genotypes (Figure 5a, left). This suggests that PV+ cell with broader AP durations may have a lower Fmaxinitial, a feature previously observed in atypical PV+ cells from other cortical areas (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020). To further investigate whether A1 PV+ cells from control mice could be functionally segregated into distinct clusters, we performed hierarchical clustering of AP half-width and Fmaxinitial values based on Euclidean distance (Figure 5a, right). This analysis identified two clusters of PV+ cells: one with short AP half-widths associated with higher Fmaxinitial values, and another with broader AP durations and lower Fmaxinitial values, consistent with what reported in other cortical regions (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020). Although hierarchical clustering distinguished these two subgroups, a few PV+ cells with longer AP half- widths exhibited Fmaxinitial values typical of PV+ cells with shorter AP half-widths (Figure 5a, right), indicating that these two parameters alone may not be sufficient to fully differentiate subtypes within our PV+ cell dataset. We thus decided to perform PCA analysis using additional key intrinsic physiological features such as passive (Vm, Rin, Cm) and active (threshold current, AP half-width, AP amplitude, first AP latency, AP threshold, fAHP amplitude, amplitude AR, frequency AR, Fmaxinitial and Fss) membrane properties (Figure 5b, left). We then selected the intersection point of the two AP half-width distributions in control and cHet mice as a cut-off to define two different subpopulations of PV+ cells: Basket Cell (BC)-short (AP half-width <0.78ms) and BC- broad (AP half-width ≥0.78ms) (Figure 5c,d). In control mice, these two PV+ cell subtypes showed major differences in Rin, Cm, Fmaxinitial and Fss (Table S1, comparison between BC-short and BC- broad in control mice). PCA analysis also revealed that, while two distinct PV+ cell subgroups were clearly distinguishable in control mice, these differences were more ambiguous in cHet mice, wherein some BC-short cells fell within the BC-broad subgroup (Figure 5d). In addition, we observed that the fraction of BC-short cells was increased in cHet compared to control mice (61% vs 41% of total PV+ cells, respectively), suggesting that Syngap1 haploinsufficiency affects specific subgroups of PV+ cells (Figure 5d).

Syngap1 haploinsufficiency in Nkx2.1+ interneurons affects the denritic arbour of a specific subpopulation of LIV PV+ cells.

Earlier version: (a, left) Strong negative correlation of Fmaxinitial with AP half-width in PV+ cells from control (blue, n=33 cells, 15 mice) and cHet mice (red, n=40 cells, 17 mice). (a, right) Hierarchical clustering based on Euclidean distance of PV+ cells from control mice. Clustering is based on AP half-width and Max frequency. Asterisks indicate cells with longer AP half-width felling into the cluster including PV+ cells with higher of values of Fmaxinitial. (b, left) Correlation of parameters describing membrane properties of PV+ interneurons. The 13 passive and active membrane properties used for PCA analysis (derived from 27 PV+ cells from control mice; see materials and methods) are arrayed against each other in a correlation matrix with the degree of correlation indicated by the shading: white is negatively correlated (correlation index of 0), black is positively correlated (correlation index of 1, diagonal squares) and light gray not correlated (correlation index of 0). PCA on the 13 parameters to reduce the dimensionality. (b, right) The first (PC1) and second (PC2) PC values derived for each interneuron are plotted against each other. No clear separation of subgroups in scatterplot of first 2 PCs is present when genotype is taken into consideration. (c) Cumulative histograms of AP half-widths in control (n=33 cells, 15 mice) and cHet mice (n=40 cells, 17 mice) fitted with two Gaussian curves. Vertical line indicates the cutoff value at intersection between the two curves. For both histograms, bins are 0.05 msec wide. (d) PCA analysis using the cutoff value of 0.78 ms and the 13 passive and active membrane properties distinguish two subgroups of PV+ cells with short (black circles) and broad (turquoise circles) AP-half width duration in both genotypes. Insets illustrate pie charts describing the % of two subgroups of PV+ cells in the control and cHet mice. (e) Anatomical reconstructions of a BC-short and (f) a BC-broad filled with biocytin in control mice during whole-cell patch-clamp recordings and post hoc immunohistochemical validation for PV. (g) Summary data in control mice (gray, BC-short n=5 cells, 4 mice; turquoise, BC-broad n=5 cells, 4 mice) showing no significant difference in terms of distance from pia (p=0.856, LMM) for both subtypes of PV+ cell analyzed indicating LIV location and significant differences in dendritic parameters between the two subpopulations of PV+ cells (LMM, *p= 0.016 for dendr. surface area, *p=0.043 for # branching points) and no change in total dendritic length (LMM, p=0.057). (h) Summary data in cHet mice (gray, BC-short n=6 cells, 4 mice; turquoise, BC-broad n=6 cells, 3 mice) showing no significant difference in terms of distance from pia (LMM, p=0.594) for both subtypes of PV+ cell and all dendritic parameters (LMM, p= 0.062 for total dendritic length, p=0.731 for dendr. surface area, p=0.081 for # branching points). (i) Summary data showing a significant increase in dendritic complexity between control (gray, n=5 cells, 4 mice) and cHet (white, n=6 cells, 4 mice) for the subpopulation of BC-short (LMM, *p=0.009 for dendr. surface area, *p=0.048 for # branching points) and no difference for the total dendritic length (LMM, p=0.070). (j) Summary data showing preserved dendritic parameters in cHet (turquoise filled with pattern, n=6 cells, 3 mice) Vs control (turquoise, n=5 cells, 4 mice) (LMM, p= 0.967 for total dendritic length, p=0.784 for dendr. surface area, p=0.290 for # branching points). (k) The strong positive correlation of dendritic surface area with AP half-width is present only in PV+ cells from control mice (blue, n=10 cells, 8 mice) and disappears in cHet mice (red, 12 cells, 7 mice).

Revised version: (a, left) Strong negative correlation of Fmaxinitial with AP half-width in PV+ cells from control (blue, n=33 cells, 15 mice) and cHet mice (red, n=40 cells, 17 mice). (a, right) Hierarchical clustering based on Euclidean distance of PV+ cells from control mice. Clustering is based on AP half-width and Max frequency. Asterisks indicate cells with longer AP half-width felling into the cluster including PV+ cells with higher of values of Fmaxinitial. (b, left) Correlation of parameters describing membrane properties of PV+ interneurons. The 13 passive and active membrane properties used for PCA analysis (derived from 27 PV+ cells from control mice; see materials and methods) are arrayed against each other in a correlation matrix with the degree of correlation indicated by the shading: white is negatively correlated (correlation index of 0) ,black is positively correlated (correlation index of 1, diagonal squares) and light gray not correlated (correlation index of 0). PCA on the 13 parameters to reduce the dimensionality. (b, right) The first (PC1) and second (PC2) PC values derived for each interneuron are plotted against each other. No clear separation of subgroups in scatterplot of first 2 PCs is present when genotype is taken into consideration. (c) Cumulative histograms of AP half-widths in control (n=33 cells, 15 mice) and cHet mice (n=40 cells, 17 mice) fitted with two Gaussian curves. Vertical line indicates the cutoff value at intersection between the two curves. For both histograms, bins are 0.05 msec wide. (d) PCA analysis using the cutoff value of 0.78 ms and the 13 passive and active membrane properties distinguish two subgroups of PV+ cells with short (black circles) and broad (turquoise circles) AP-half width duration in both genotypes. Insets illustrate pie charts describing the % of two subgroups of PV+ cells in the control and cHet mice. (e) Anatomical reconstructions of a BC- short and (f) a BC-broad filled with biocytin in control mice during whole-cell patch-clamp recordings and post hoc immunohistochemical validation for PV. (g) Summary data in control mice (gray, BC-short n=5 cells, 4 mice; turquoise, BC-broad n=5 cells, 4 mice) showing no significant difference in terms of distance from pia (p=0.856, LMM) for both subtypes of PV+ cell analyzed indicating LIV location and significant differences in dendritic parameters between the two subpopulations of PV+ cells (LMM , *p= 0.016 for dendr. surface area, *p=0.043 for # branching points) and no change in total dendritic length (LMM , p=0.057). (h) Summary data in cHet mice (gray, BC-short n=6 cells, 4 mice; turquoise, BC-broad n=6 cells, 3 mice) showing no significant difference in terms of distance from pia (LMM, p=0.594) for both subtypes of PV+ cell and all dendritic parameters (LMM, p= 0.062 for total dendritic length, p=0.731 for dendr. surface area, p=0.081 for # branching points). (i) Summary data showing a significant increase in dendritic complexity between control (gray, n=5 cells, 4 mice) and cHet (white, n=6 cells, 4 mice) for the subpopulation of BC-short (LMM , *p=0.009 for dendr. surface area, *p=0.048 for # branching points) and no difference for the total dendritic length (LMM, p=0.070). (j) Summary data showing preserved dendritic parameters in cHet (turquoise filled with pattern, n=6 cells, 3 mice) Vs control (turquoise, n=5 cells, 4 mice) (LMM, p= 0.967 for total dendritic length, p=0.784 for dendr. surface area, p=0.290 for # branching points). (k) The strong positive correlation of dendritic surface area with AP half-width is present only in PV+ cells from control mice (blue, n=10 cells, 8 mice) and disappears in cHet mice (red, 12 cells, 7 mice). * indicates p value <0.05.

Earlier version: Next, we examined whether the diversity in PV+ cell electrophysiological profiles was reflected in their dendritic arborisation (Figure 6e-h, Table 6.1, 6.2, 6.3, 6.4). We also quantified somata anatomical location (distance in µ from pia) to confirm that the recorded and analysed cells were located in LIV (Figure 6g,h). In control mice, both BC-short and BC-broad cells showed ovoid somata and multipolar dendrites (Figure 6e,f). Anatomical reconstruction and morphometric analysis revealed differences in dendritic arborization that correlated positively with AP half-width in control mice (Figure 6g,k). In particular, BC-short cells showed significantly lower branch point numbers and dendritic surface area as compared to BC-broad cells (Figure 6g). In contrast, the dendritic arbor of BC-short neurons vs BC-broad did not show significant differences in cHet mice (Figure 6h). Direct comparison of PV+ cell dendritic arbor in cHet vs control littermates clearly showed that BC-short neurons were specifically affected by Syngap1 haploinsufficiency (Figure 6i, j), with BC-short cells from cHet mice showing a significant increase in dendritic complexity compared to those from control mice (Figure 6i). Further, the strong positive correlation of dendritic surface area with AP half-width was present only in PV+ cells from control mice but disappeared in cHet mice (Figure 6k). Altogether, these data indicate that embryonic-onset Syngap1 haploinsufficiency in Nkx2.1+ interneurons alters the dendritic development of a specific PV+ cell subpopulation, which could in turn affect their intrinsic excitability and dendritic integration of synaptic inputs.

Revised version: Next, we examined whether the diversity in PV+ cell electrophysiological profiles was reflected in their dendritic arborisation (Figure 5e-h, Table 5.1-5.4). We also measured the anatomical location of cell bodies (distance in µm from the pia) to confirm that the recorded and analyzed cells were situated within LIV (Figure 5g,h). In control mice, both BC-short and BC-broad cells showed ovoid somata and multipolar dendrites (Figure 5e,f). Anatomical reconstruction and morphometric analysis revealed differences in dendritic arborization that correlated positively with AP half-width in control mice (Figure 5g,k). In particular, BC-short cells showed significantly lower branch point numbers and dendritic surface area as compared to BC-broad cells (Figure 5g). In contrast, the dendritic arbor of BC-short neurons vs BC-broad did not show significant differences in cHet mice (Figure 5h). Direct comparison of PV+ cell dendritic arbor in cHet vs control littermates clearly showed that BC-short neurons were specifically affected by Syngap1 haploinsufficiency (Figure 5i, j), with BC-short cells from cHet mice showing a significant increase in dendritic complexity compared to those from control mice (Figure 5i). Further, the strong positive correlation of dendritic surface area with AP half-width was present only in PV+ cells from control mice but disappeared in cHet mice (Figure 5k). Altogether, these data indicate that embryonic-onset Syngap1 haploinsufficiency in Nkx2.1+ interneurons alters dendritic development in a specific subpopulation of PV+ cells, leading to the increased dendritic area and complexity. These structural changes may, in turn, affect their intrinsic excitability and the dendritic integration of synaptic inputs.

Earlier version: Based on the observed heterogeneity in morpho-electric parameters of PV+ cells, we next sought to investigate the effect of Syngap1 haploinsufficiency on the intrinsic excitability of specific PV+ cell subtypes (Figure 7, Tables 7.1, 7.2). We found that BC-short cells showed preserved passive membrane properties (Fig. 7a) and altered active membrane properties (Figure 7b-e) in cHet compared to control mice. In particular, we found increased AP threshold affecting AP amplitude (Figure 7b,c) and increased rheobase current (Figure 7d), indicating a decrease in the excitability of cHet BC-short cells. cHet BC-short interneurons displayed AP patterns similar to those in control BC-short (Figure 7e, right), but fired less APs in response to somatic depolarization (Figure 7e, left). In contrast, BC-broad neurons had a more hyperpolarized RMP (Figure 7f), and increased AP latency and threshold (Figure 7g,h) in cHet mice compared to controls. However, in cHet BC-broad neurons these changes were not translated into decreased ability to generate spikes (Figure 7i,j). Altogether these data suggest that BC-short neurons may be overall more vulnerable to Syngap1 haploinsufficiency than BC-broad neurons. They further indicate that Syngap1 levels appears to play a common role in determining the threshold for AP generation in all adult PV+ cells.

Revised version: Based on the observed heterogeneity in morpho-electric parameters of PV+ cells, we next sought to investigate the effect of Syngap1 haploinsufficiency on the intrinsic excitability of specific PV+ cell subtypes (Figure 6, Tables 6.1, 6.2). We found that BC-short cells showed preserved passive membrane properties (Figure 6a) but altered active membrane properties (Figure 6b-e) in cHet compared to control mice. In particular, we found increased AP threshold affecting AP amplitude (Figure 6b,c) and increased threshold current (Figure 6d), indicating a decrease in the excitability of cHet BC-short cells. cHet BC-short interneurons displayed AP firing patterns similar to those in control BC-short (Figure 6e, right), but fired less APs in response to somatic depolarization (Figure 6e, left, Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ****p<0.0001). In contrast, BC-broad neurons had a more hyperpolarized RMP (Figure 6f), and increased AP latency and threshold (Figure 6g,h) in cHet mice compared to controls. However, in cHet BC-broad neurons these changes were not translated into decreased ability to generate spikes (Figure 6i,j, Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, p= 0.333). Altogether these data suggest that BC-short neurons may be overall more vulnerable to Syngap1 haploinsufficiency than BC-broad neurons. They further indicate that Syngap1 levels appears to play a common role in determining the threshold for AP generation in all adult PV+ cells.

Intrinsic excitability is decreased in both subpopulations of PV+ cells in cHet mice

Earlier version: (a) Summary data showing no changes in the passive membrane properties of BC-short between control (blue, n=12 cells, 9 mice) and cHet mice (red, n=24 cells, 13 mice) (LMM, p=0.189 for Vm, p=0.856 for Rin, p=0.188 for Cm, p=0.077 for Tm) (b) Summary data showing no differences in AP half-width (p=0.386, LMM) and AP latency (LMM, p=0.210) but a significant decrease in AP amplitude (LMM, *p=0.024) of BC-short recorded in cHet mice. (c, left) Summary bar graph shows a significant decrease in AP threshold from cHet mice (LMM, **p=0.002) and phase plane (c, center) for the first AP generated confirmed this result with a more hyperpolarized value (the area squared) of AP threshold for generation of AP. (c, right) Representative single APs evoked by rheobase currents from control and cHet mice. (d) Summary bar graph shows a significant increase in the rheobase current (LMM,*p=0.015). (e, left) Summary plot showing a reduction of averaging number of APs per current step (40 pA) amplitude recorded from LIV BC-short in cHet (red circles, n = 22 cells; 13 mice) compared to control (blue circles, n = 11 cells, 9 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, **p=0.005). (e, right) Representative voltage responses indicating the typical FS firing pattern of BC-short in control and cHet mice in response to depolarizing (+120 pA and +240 pA) current injections corresponding to rheobase and 2x rheobase current. (f) Summary data showing a significant decrease in RMP (LMM, *p=0.023) of BC-broad from cHet mice (red, n=16 cells, 11 mice) but no changes in the other passive membrane properties compared to control mice (blue, n=21 cells, 12 mice) (LMM, p=0.244 for Rin, p=0.170 for Cm, p=0.639 for Tm). (g) Summary data showing no differences in AP half-width (LMM, p=0.593) and AP amplitude (LMM, p=0.713) and a significant increase in AP latency (LMM, *p=0.035) from BC-broad cells recorded in cHet mice. (h, left) Summary bar graph shows a significant decrease in AP threshold from cHet mice (LMM, *p=0.010) and phase plane (h, center) for the first AP generated confirmed this result with a more hyperpolarized value (the area squared) of AP threshold for generation of AP. (h, right). Representative single APs evoked by rheobase currents from control and cHet mice. (i) Summary bar graph shows no difference in the rheobase current (LMM, p=0.402). (j, left) Summary plot showing no difference in the averaging number of APs per current step (40 pA) amplitude recorded from LIV BC-broad in cHet (red circles, n=16 cells, 11 mice) compared to control (blue circles, n=18 cells; 11 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, p= 0.333). (j, right) Representative voltage responses indicating the typical FS firing pattern of BC broad in control and cHet mice in response to depolarizing (+120 pA and +240 pA) current injections corresponding to rheobase and 2x rheobase current.

Revised version: (a) Summary data showing no changes in the passive membrane properties of BC-short between control (blue, n=12 cells, 9 mice) and cHet mice (red, n=24 cells, 13 mice) (LMM , p=0.189 for Vm, p=0.856 for Rin, p=0.188 for Cm, p=0.077 for τ) (b) Summary data showing no differences in AP half-width (p=0.386, LMM) and AP latency (LMM, p=0.210) but a significant decrease in AP amplitude (LMM, *p=0.024) of BC-short recorded in cHet mice. (c, left) Summary bar graph shows a significant increase in AP threshold from cHet mice (LMM, **p=0.002). (c, right) Representative single APs evoked by threshold currents from control and cHet mice. (d) Summary bar graph shows a significant increase in the threshold current (LMM, *p=0.015). (e, left) Summary plot showing a reduction of averaging number of APs per current step (40 pA) amplitude recorded from LIV BC- short in cHet (red circles, n = 22 cells; 13 mice) compared to control (blue circles, n = 11 cells, 9 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, **** p<0.0001). (e, right) Representative voltage responses indicating the typical FS firing pattern of BC-short in control and cHet mice in response to depolarizing (+120 pA and +240 pA) current injections corresponding to threshold current and 2x threshold current. (f) Summary data showing a significant decrease in RMP (LMM, *p=0.023) of BC-broad from cHet mice (red, n=16 cells, 11 mice) but no changes in the other passive membrane properties compared to control mice (blue, n=21 cells, 12 mice) (LMM, p=0.244 for Rin, p=0.170 for Cm, p=0.639 for τ). (g) Summary data showing no differences in AP half-width (LMM, p=0.593) and AP amplitude (LMM, p=0.713) and a significant increase in AP latency (LMM, *p=0.035) from BC-broad cells recorded in cHet mice. (h, left) Summary bar graph shows a significant increase in AP threshold from cHet mice (LMM, *p=0.010). (h, right) Representative single APs evoked by threshold currents from control and cHet mice. (i) Summary bar graph shows no difference in the threshold current (LMM, p=0.402). (j, left) Summary plot showing no difference in the averaging number of APs per current step (40 pA) amplitude recorded from LIV BC-broad in cHet (red circles, n=16 cells, 11 mice) compared to control (blue circles, n=18 cells; 11 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, p= 0.333). (j, right) Representative voltage responses indicating the typical FS firing pattern of BC broad in control and cHet mice in response to depolarizing (+120 pA and +240 pA) current injections corresponding to threshold current and 2x threshold current. * indicates p value <0.05; ** indicates p value <0.005.

Added in the revised version: Syngap1 haploinsufficiency in MGE-derived interneurons affects SST+ interneurons firing and spontaneous excitatory inputs

Removed in the revised version: SST+ interneuron intrinsic and firing properties are less affected by Syngap1 haploinsufficiency in MGE-derived interneurons as compared to PV+ cells

Earlier version: Next, we examined whether Syngap1 haploinsufficiency in MGE-derived interneurons affects the other major group of Nkx2.1-expressing cortical interneurons, the SST+ interneurons (Figure 8 a-f, Table 8). In current-clamp recordings, control SST+ cells displayed a low discharge rate and typical AP frequency accommodation in response to incremental steps of current injection (Figure 8a). The molecular identity of this interneuron subtype was then confirmed by immunopositivity for SST+ and immunonegativity for PV (Figure 8b). Interestingly, PCA analysis using the previously mentioned electrophysiological parameters clearly distinguished SST+ neurons from BC-short subtype of PV+ cells, but showed an overlap between BC-broad PV+ and SST+ cells (Figure 8c). These data indicate that, in mature A1, a subtype of PV+ cells share some electrophysiological features with SST+ cells, indicating the necessity to perform post-hoc immunohistochemical validation (Figure 8b,c). cHet SST+ cells showed no significant changes in active and passive membrane properties (Figure 8d,e); however, their evoked firing properties were affected with fewer APs generated in response to the same depolarizing current injection compared to control SST+ cells (Figure 8f).

Revised version: To investigate whether the observed alterations in the synaptic and intrinsic properties of PV+ interneurons are cell-type-specific, we next examined whether the second major group of Nkx2.1- expressing cortical interneurons, the SST+ interneurons, exhibits similar changes. In current-clamp recordings, control SST+ cells displayed a low firing rate and characteristic AP frequency accommodation in response to incremental current injections (Figure 7a, Table 7). Morphologically, SST neurons showed ovoid-shaped somata, multipolar (Suppl. Figure 1b, left) or bitufted (Suppl. Figure 1b, right) dendritic arbours with spines. Their axon projected into layer I, where it arborized giving rise to multiple collaterals. The molecular identity of this interneuron subtype was confirmed by immunopositivity for SST+ and immunonegativity for PV (Figure 7b, Suppl. Figure 1b). Interestingly, PCA analysis using the previously mentioned electrophysiological parameters clearly distinguished SST+ neurons from BC-short subtype of PV+ cells, but showed an overlap between BC-broad PV+ and SST+ cells (Figure 7c). These data indicate that, in mature A1, a subtype of PV+ cells share some electrophysiological features with SST+ cells, indicating the necessity to perform post-hoc immunohistochemical validation (Figure 7b,c). cHet SST+ cells showed no significant changes in the active or passive membrane properties we analysed (Figure 7d,e). However, their evoked firing properties were affected, with fewer APs in response to the same depolarizing current injection compared to control SST+ cells (Figure 7f, Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ***p<0.001). Furthermore, we found that sEPSC amplitude and charge transfer were significantly decreased in SST+ neurons from cHet mice compared to control littermates (Suppl. Figure 4; Table S2).

Earlier version: Intrinsic excitability of SST+ cells is less affected by embryonic-onset Syngap1 haploinsufficiency in Nkx2.1 interneurons.

Revised version: Evoked firing properties are reduced in SST+ cells from mice with embryonic-onset Syngap1 haploinsufficiency in Nkx2.1 interneurons.

Earlier version: (a) Representative voltage responses indicating the typical regular adapting firing pattern of SST+ in control mice in response to hyperpolarizing (−40 pA) and depolarizing (+80 pA and +160 pA) current injections corresponding to Ih associated voltage rectification, rheobase and 2x rheobase current respectively. (b) Post hoc immunohistochemical validation of these interneurons confirming their positivity for SST+ and negativity for PV-. (c) PCA using the 13 parameters previously described clearly separate the cluster of SST+ cells (pink circles) from BC-short (black circles) having however some overlaps with BC-broad (turquoise circles) in control mice. (d) Summary data showing no changes in the passive (LMM, p=0.469 for Vm, p=0.681 for Rin, p=0.922 for Cm, p=0.922 for Tm) and (e) active membrane properties (LMM, p=0.675 for AP half-width, p=0.342 for AP amplitude, p=0.081 for AP latency, p=0.119 for AP threshold) between SST+ cells from control (blue, n=10 cells, 7 mice) and cHet mice (red, n=12 cells, 7 mice) (f) Summary plot showing a reduction of averaging number of APs per current step (40 pA) amplitude recorded from LIV SST+ in cHet (red circles, n= 12 cells, 8 mice) compared to control (blue circles, n=8 cells, 7 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ****p<0.0001).

Revised version: (a) Representative voltage responses indicating the typical regular adapting firing pattern of SST+ in control mice in response to hyperpolarizing (-40 pA) and depolarizing (+80 pA and +160 pA) current injections corresponding to Ih associated voltage rectification, threshold current and 2x threshold current respectively. (b) Post hoc immunohistochemical validation of these interneurons confirming their positivity for SST+ and negativity for PV-. (c) PCA using the 13 parameters previously described clearly separate the cluster of SST+ cells (pink circles) from BC-short (black circles) having however some overlaps with BC-broad (turquoise circles) in control mice. (d) Summary data showing no changes in the passive (LMM, p=0.283 for Vm, p=0.959 for Rin, p=0.484 for Cm, p=0.501 for τ) and (e) active membrane properties (LMM, p=0.332 for AP half-width, p=0.126 for AP amplitude, p=0.296 for AP latency, p=0.154 for AP threshold) between SST+ cells from control (blue, n=11 cells, 8 mice) and cHet mice (red, n=16 cells, 9 mice) (f) Summary plot showing a reduction of averaging number of APs per current step (40 pA) amplitude recorded from LIV SST+ in cHet (red circles, n= 16 cells, 9 mice) compared to control (blue circles, n=11 cells, 8 mice) neurons (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, ***p<0.001). * indicates p value <0.05; ** indicates p value <0.005.

Added in the revised version: Thus, embryonic-onset Syngap1 haploinsufficiency in MGE-derived interneurons significantly reduced spontaneous excitatory inputs and evoked firing properties of both PV+ and SST+. However, alterations in AP initiation, as indicated by an increased firing were observed exclusively in PV+ cells (Figure 6c,h), suggesting a selective impairment in their intrinsic excitability.

Removed in the revised version: Thus, embryonic-onset Syngap1 haploinsufficiency in MGE-derived interneurons substantially reduced the excitability of auditory cortex LIV PV+ and SST cells, with a major impact on the PV+ cell population as reflected in both single AP properties and AP firing pattern.

A selective Kv1-blocker rescues PV+ cell intrinsic excitability in cHet mice

Earlier version: Based on previous studies performed in PV+ cells (Wang et al., 1994; Goldberg et al., 2008; Zurita et al., 2018), changes in voltage-gated D-type K+ currents mediated by the Kv1 subfamily could account for the observed altered intrinsic excitability observed in cHet mice. To test this hypothesis, we performed current-clamp recordings from PV+ cells in control and cHet mice in presence or absence of α-DTX (100 nM), a specific blocker of channels containing Kv1.1, Kv1.2 or Kv1.6 (Figure 9a-e, Table 9). As hypothesized, the presence of α-DTX rescued the voltage threshold for AP generation in cHet PV+ cells to control levels, without affecting the AP shape parameters (Figure 9 a,b,c, Table 9), while the relation between the AP number and current injection remained the same (Figure 9d), indicating that α-DTX had no impact on PV+ cell firing. Further, α-DTX facilitated AP initiation in cHet PV+ cells by reducing AP delay from stimulation onset (Figure 9e). These data revealed that Syngap1 haploinsufficiency potentially affects voltage-gated D-type K+ currents, thereby decreasing the excitability of PV+ cells in adult A1.

Revised version: Based on previous studies performed in PV+ cells (Wang et al., 1994; Goldberg et al., 2008; Zurita et al., 2018), changes in voltage-gated D-type K+ currents mediated by the Kv1 subfamily could account for the observed altered AP threshold observed in cHet mice. To test this hypothesis, we performed current-clamp recordings from PV+ cells in control and cHet mice in presence or absence of α-DTX (100 nM), a specific blocker of channels containing Kv1.1, Kv1.2 or Kv1.6 (Figure 8a-e, Table 8). As hypothesized, the presence of α-DTX rescued the voltage threshold for AP generation in cHet PV+ cells to control levels, without affecting the AP shape (Figure 8 a,b,c, Table 8), while the relation between the AP number and current injection remained the same (Figure 8d), indicating that α-DTX had no impact on PV+ cell firing. Further, α-DTX facilitated AP initiation in cHet PV+ cells by reducing AP delay from stimulation onset (Figure 8e). These data suggest that Syngap1 haploinsufficiency may enhance voltage-gated D-type K+ currents, thereby leading to reduced excitability of PV+ cells in the adult A1.

Syngap1 haploinsufficiency alters the intrinsic excitability of LIV PV+ cells by affecting voltage-gated D-type K+ currents.

Earlier version: (a, left) Summary bar graph shows a significant decrease in AP half-width in PV+ cells from cHet (red) vs control (blue) mice (LMM, *p=0.034), which persist when cHet PV+ cells are treated with α-DTX (red with diagonal stripes, LMM, *p=0.039). (a, right) Summary bar graph shows a significant decrease in AP threshold of PV+ cells from vehicle-treated cHet mice (red) compared to vehicle-treated control mice (blue, LMM, *p=0.049) and the rescue of this deficit in presence of α-DTX (blue vs red with diagonal stripes, LMM, p=0.940). (b) Delta (Δ) value was calculated for AP threshold by subtracting individual values of α-DTX-treated cells from the average of their respective control group. A significant increase in AP threshold Δ number was found for cHet α-DTX-treated PV+ cells compared to control α-DTX-treated PV+ cells (LMM, *p=0.015). (c. left) Phase plane for the first AP generated confirmed this result with more depolarized value (the area squared) of AP threshold in cHet α-DTX-treated cells (red dotted line). (c, center and right) Representative single APs evoked by rheobase currents from vehicle-treated control (blue) and cHet (red) mice (center), and control (blue dotted line) and cHet α-DTX-treated (red dotted line) PV+ cells. (d) Summary plot showing no difference in the averaging number of APs per current step (40 pA) amplitude recorded from LIV PV+ in cHet and control, both α-DTX-treated and vehicle-treated, PV+ cells (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, p>0.05). (e, left) Summary bar graph shows a significant difference in AP latency Δ number in α-DTX-treated cHet vs α-DTX-treated control PV+ cells (LMM, *p=0.006). (e, right) Representative voltage traces clearly show a reduction in the AP onset for cHet PV+ cells treated with α-DTX (pink trace) compared to vehicle-treated cHet PV+ cells (red trace), while control PV+ cells are affected (vehicle treated-control PV+cells, blue traces; control α-DTX PV+ cells, light blue traces). Control mice: vehicle treated, n=9 cells from 4 mice; α-DTX-treated, n=11 cells from 6 mice; cHet mice: vehicle treated, n=23 cells from 10 mice; cHet α-DTX-treated, n=18 cells from 8 mice.

Revised version: (a, left) Summary bar graph shows a significant decrease in AP half-width in PV+ cells from cHet (red) vs control (blue) mice (LMM, *p=0.034), which persist when cHet PV+ cells are treated with α-DTX (red with diagonal stripes, LMM, *p=0.039). (a, right) Summary bar graph shows a significant increase in AP threshold of PV+ cells from vehicle-treated cHet mice (red) compared to vehicle-treated control mice (blue, LMM, *p=0.049) and the rescue of this deficit in presence of α-DTX (blue vs red with diagonal stripes, LMM, p=0.940). (b) Delta (Δ) value was calculated for AP threshold by subtracting individual values of α-DTX-treated cells from the average of their respective control group. A significant increase in AP threshold Δ number was found for cHet α-DTX-treated PV+ cells compared to control α-DTX-treated PV+ cells (LMM, *p=0.015). (c) Representative single APs evoked by threshold currents from vehicle-treated control (blue) and cHet (red) mice (center), and control (blue dotted line) and cHet α-DTX-treated (red dotted line) PV+ cells. (d) Summary plot showing no difference in the averaging number of APs per current step (40 pA) amplitude recorded from LIV PV+ in cHet and control, both α-DTX-treated and vehicle-treated, PV+ cells (Two-way Repeated Measure ANOVA with Sidak’s multiple comparison post hoc test, p>0.05). (e, left) Summary bar graph shows a significant difference in AP latency Δ number in α-DTX-treated cHet vs α-DTX-treated control PV+ cells (LMM, *p=0.006). (e, right) Representative voltage traces clearly show a reduction in the AP onset for cHet PV+ cells treated with α-DTX (pink trace) compared to vehicle-treated cHet PV+ cells (red trace), while control PV+ cells are not affected (vehicle treated-control PV+cells, blue traces; control α-DTX PV+ cells, light blue traces). Control mice: vehicle treated, n=9 cells from 4 mice; α-DTX-treated, n=11 cells from 6 mice; cHet mice: vehicle treated, n=23 cells from 10 mice; cHet α-DTX-treated, n=18 cells from 8 mice. * indicates p value <0.05.

Discussion

Earlier version: In this study, we tested whether and how intrinsic and synaptic properties of two main cortical GABAergic subtypes, PV+ and SST+ interneurons, are affected by embryonic-onset MGE-restricted Syngap1 haploinsufficient mice. Our data demonstrated that, at least in A1, PV+ cells are particularly vulnerable to early-onset Syngap1 haploinsufficiency. In particular, mutant PV+ cells showed reduced intracortical and thlamo-cortical synaptic drive, contrary to Syngap1’s documented role in the formation and plasticity of glutamatergic synapses on excitatory cells. We further found that Syngap1 haploinsufficiency alters PV+, but not SST+, cell intrinsic properties, overall leading to decreased intrinsic excitability affecting its input-output function. PV+ cell intrinsic excitability was rescued by pharmacological inhibition of voltage-gated D-type K+ currents mediated by the Kv1 subfamily, suggesting their potential involvement as molecular mediators of functional deficits induced by Syngap1 haploinsufficiency. Syngap1 has been studied mainly in the context of synaptic physiology; therefore, our data highlights a novel aspect of Syngap1 biology. Since Syngap1 mRNA expression in PV+ and SST+ cells is not limited to A1 (Zhao and Kwon, 2023; Jadhav et al, 2024), it is likely that its haploinsufficiency may affect interneurons physiology in other cortical regions, as well. In our mouse model, Syngap1 haploinsufficiency is driven by the expression of Nkx2.1, which has an embryonic onset (E10.5). However, Syngap1 expression is thought to be highest during periods of robust synaptogenesis (Porter et al., 2005; McMahon et al., 2012; Gou et al., 2020; Jadhav et al, 2024). The precise developmental time window during which Syngap1 insufficiency disrupt PV+ neuron properties remain thus to be determined.

Revised version: In this study, we investigated the impact of embryonic-onset, MGE-restricted Syngap1 haploinsufficiency on the intrinsic and synaptic properties of the two major cortical GABAergic interneuron subtypes, PV+ and SST+ cells. In particular, mutant PV+ cells showed reduced intracortical and thalamo-cortical glutamatergic synaptic drive. We further found that Syngap1 haploinsufficiency has a significant impact on the intrinsic properties, in particular AP threshold, of PV+ cells, resulting in overall decreased excitability. The intrinsic excitability of PV+ cells was rescued in part by pharmacological inhibition of voltage-gated D-type K+ current mediated by the Kv1 subfamily, suggesting that this current may serve as molecular mediator of the functional deficits induced by Syngap1 haploinsufficiency. Syngap1 has been studied mainly in the context of synaptic physiology; therefore, our data highlights a novel aspect of Syngap1 biology.

Added in the revised version: Since Syngap1 mRNA expression in PV+ and SST+ cells is not limited to A1 (Zhao and Kwon, 2023; Jadhav et al, 2024), it is likely that its haploinsufficiency may affect interneurons physiology in other cortical regions, as well. In our mouse model, Syngap1 haploinsufficiency is driven by the expression of Nkx2.1, which has an embryonic onset (E10.5). However, Syngap1 expression is thought to be highest during critical periods of synaptogenesis (Porter et al., 2005; McMahon et al., 2012; Gou et al., 2020; Jadhav et al, 2024). Therefore, the precise developmental time window during which Syngap1 insufficiency disrupts PV+ neuron properties, remains to be determined.

Added in the revised version: Syngap1 is thought to be a potent regulator of excitatory synapses and its reduced expression in excitatory cells causes an increase in AMPA receptor density and premature maturation of excitatory synapses (Rumbaugh et al., 2006; Clement et al., 2012; 2013). Unexpectedly, here we found that Syngap1 haploinsufficiency restricted to MGE-derived interneurons depresses glutamatergic synaptic transmission, potentially via presynaptic mechanisms. In addition, LIV PV+ cells receive the strongest thalamocortical input compared to excitatory cells and other subpopulation of GABAergic interneurons (Ji et al., 2016; Zuritasc et al., 2018). Our study indicates that a reduction in AMPA-mediated thalamocortical transmission onto LIV PV+ cells may contribute to the deficit in glutamatergic drive, as suggested by facilitated PPR in cHet compared to control mice. This, combined with the increased onset latencies of thalamocortical-evoked AMPA responses along with an enhanced NMDA/AMPA ratio, points to both a likely decrease in presynaptic release from thalamocortical fibers and impaired recruitment of PV+ cells by thalamic inputs in cHet mice.

Earlier version: Syngap1 is thought to be a potent regulator of excitatory synapses and its reduced expression in excitatory cells causes an increase in AMPA receptor density and premature maturation of excitatory synapses (Rumbaugh et al., 2006; Clement et al., 2012; 2013). Unexpectedly, here we found that Syngap1 haploinsufficiency restricted to MGE-derived interneurons depresses AMPA-mediated synaptic transmission, potentially with the contribution of a presynaptic mechanism involving the reduction of vGlut1+ glutamatergic boutons. In addition, LIV PV+ cells receive the strongest thalamocortical input compared to excitatory cells and other subpopulation of GABAergic interneurons (Ji et al., 2016; Zurita et al., 2018). Our study showed that one of the sources of the deficit in glutamatergic drive may arise from a decrease in the AMPA-mediated thalamocortical transmission onto LIV PV+ cells, as indicated by facilitated PPR in cHet compared to control mice, suggesting a decrease in presynaptic release from thalamocortical fiber in cHet mice. Interestingly, we observed increased onset latencies of thalamocortical evoked AMPA together with enhanced NMDA/AMPA ratio, indicating a kinetically slower thalamic recruitment of PV+ cells. How could Syngap1 haploinsufficiency in Nkx2.1-expressing cells affect the glutamatergic drive coming from local and thalamic excitatory cells? Our data suggest a role of Syngap1 in promoting GABAergic cell intrinsic excitability, particularly in cortical PV+ interneurons. These findings are in line with recent data reporting a decrease in intrinsic excitability of developing cortical excitatory cells (Arora et al., 2022), and highlight the emerging function of Syngap1 in the somato-dendritic compartment (Arora et al., 2022). Since in PV+ cells, connectivity and cell excitability are reciprocally regulated at the circuit level (Favuzzi et al., 2017), it is possible that early-onset Syngap1 haploinsufficiency in MGE-derived interneurons may first affect the development of their intrinsic excitability properties, which in turn would modulate the maturation of their excitatory drive, and activity levels in adulthood. Alternatively, homeostatic adaptation of PV+ interneurons in response to the decreased number of excitatory inputs could trigger changes in voltage-gated D-type K+ currents (Dehorter et al., 2015; Favuzzi et al., 2017). Therefore, Syngap1 may play complex roles at the cellular level and at different developmental stages, based on the crosstalk between neuronal activity levels and Syngap1 localisation and interaction with other proteins. The use of a conditional genetic strategy to induce Syngap1 haploinsufficiency specifically in MGE-derived interneurons allow investigating its effects in these GABAergic populations; however, it is important to highlight the limit of this approach, since whether GABAergic interneurons physiology would be similarly affected within an entire network carrying the same mutation (thus affecting also excitatory neurons) remains to be established.

Revised version: How could Syngap1 haploinsufficiency in Nkx2.1-expressing cells affect the glutamatergic drive coming from local and thalamic excitatory cells? Our data suggest a role of Syngap1 in promoting GABAergic cell intrinsic excitability during normal development. These findings are in line with recent data reporting a decrease in intrinsic excitability of developing cortical excitatory cells in Syngap1+/- mice (Arora et al., 2022). In PV+ cells, connectivity and cell excitability are reciprocally regulated at the circuit level (Favuzzi et al., 2017); thus, it is possible that early-onset Syngap1 haploinsufficiency in MGE-derived interneurons may first affect the development of their intrinsic excitability properties, which in turn would modulate the maturation of their excitatory drive, and activity levels in adulthood. Alternatively, homeostatic adaptation of PV+ interneurons in response to the decreased number of excitatory inputs could trigger changes in voltage-gated D-type K+ currents (Dehorter et al., 2015; Favuzzi et al., 2017). Therefore, Syngap1 may play complex roles at the cellular level and at different developmental stages, based on the crosstalk between neuronal activity levels and Syngap1 localisation and interaction with other proteins. The use of a conditional genetic strategy to induce Syngap1 haploinsufficiency specifically in MGE-derived interneurons allow investigating its effects in these GABAergic populations; however, it is important to highlight the limit of this approach, since whether GABAergic interneurons physiology would be similarly affected within an entire network carrying the same mutation (thus affecting also excitatory neurons) remains to be established.

Earlier version: Based on morphology and synaptic targets, 3 main subgroups of PV+ cells have been identified in auditory cortex, ie basket cells (BCs), chandelier cells and long-range projecting PV+ cells (Levy and Reyes, 2012, Rock et al., 2018; Zurita et al., 2018; Bertero et al., 2019). Despite differences in morphology and target selectivity, different PV+ cell subgroups share common electrophysiological features such as short AP-half width, low input resistance, very high rheobase, and relatively small AP amplitude. In particular, while clusters of atypical BC PV+ have been previously reported in several areas (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020), auditory cortex PV+ BCs have been considered as a homogeneous electrophysiological group (Studer and Barkat, 2022). One of the main findings of this study is that mature A1 contains at least 2 morphological and electrophysiological subgroups of PV+ BCs, including a subgroup with a surprising broader duration in AP half-width (Table S2). It is possible that the genetic tools used to identify PV+ cells (PV_Cre or G42 mice vs Nkx2.1_Cre mice) might enrich for a specific PV+ BC cell subtype. Alternatively, the proportion of different PV+ BCs might depend on the specific cortical region and layer. The presence of at least 2 different PV+ BC cell subtypes with different electrophysiological signatures may in part explain the previously reported variability in the recruitment of auditory cortex PV+ cells in vivo (Seybold et al., 2015; Phillips and Hasenstaub, 2016; Keller et al., 2018; Gothner et al., 2021), since differences might arise from the duration of AP half-width used to sort FS cells.

Revised version: Based on morphology and synaptic targets, three main subgroups of PV+ cells have been identified in auditory cortex, ie basket cells (BCs), chandelier cells and long-range projecting PV+ cells (Levy and Reyes, 2012, Rock et al., 2018; Zurita et al., 2018; Bertero et al., 2019). Despite differences in morphology and target selectivity, different PV+ cell subgroups share common electrophysiological features such as short AP-half width, low input resistance, very high threshold current, and relatively small AP amplitude. In particular, while clusters of atypical BC PV+ have been previously reported in several areas (Helm et al., 2013; Nassar et al., 2015; Bengtsson Gonzales et al., 2018; Ekins et al., 2020), auditory cortex PV+ BCs have been considered as a homogeneous electrophysiological group (Studer and Barkat, 2022). One of the key findings of this study is that mature A1 contains at least two distinct morphological and electrophysiological subgroups of PV+ BCs, including a subgroup with an unexpectedly broader AP half-width (Table S1). The genetic tools used to identify PV+ cells (PV_Cre or G42 mice vs Nkx2.1_Cre mice) might selectively label a specific PV+ BC cell subtype. Alternatively, the distribution of different PV+ BC subtypes could depend on the cortical region and layer. The presence of at least two PV+ BC subtypes with different electrophysiological characteristics may partly explain the previously observed variability in the recruitment of auditory cortex PV+ cells in vivo (Seybold et al., 2015; Phillips and Hasenstaub, 2016; Keller et al., 2018; Gothner et al., 2021), as these differences could arise from the AP half-width criteria used to sort FS cells.

Earlier version: Despite the differences observed in how Syngap1 haploinsufficiency affects the anatomy and physiological properties of the 2 PV+ cell subtypes, a shared deficit we observed was the decrease in intrinsic excitability, as suggested by the increased AP threshold affecting AP initiation. In PV+ cells, the voltage-gated D-type K+ currents mediated by the Kv1 family strongly contribute to AP generation, making them powerful targets to modify AP latency, threshold and rheobase current in these cells (Wang et al., 1994; Goldberg et al., 2008; Zurita et al., 2018). Here, we indeed rescued the excitability of mutant PV+ cells by pharmacologically inhibiting this channel family using α-DTX; however, whether Kv1 currents or/and channel density are increased in mutant PV+ cells remain to be investigated. Consistent with our findings, Arora et al. (2022) rescued pyramidal cell intrinsic excitability and neuronal morphology via lowering elevated potassium channels, by expressing a dominant negative form of the Kv4.2 potassium channel subunit, dnKv4.2, in developing Syngap1 mutant mice. Kv4.2 and Kv1 potassium channels modulate the intrinsic excitability of pyramidal cell and PV+ interneurons, respectively, indicating that the action of Syngap1 on potassium channels may be a general mechanism. A recent study focusing on the PSD interactomes of Syngap1 isolated from adult homogenized mouse cortex suggested a physical interaction between Syngap1 and the potassium channel auxiliary subunit Kvβ2 (Wilkinson et al., 2017). Since Kvβ2 regulates the translocation of Kv1 channels in dopaminergic neurons and potentially PV+ cells (Okaty et al., 2009; Yee et al., 2022), it is possible that Syngap1 haploinsufficiency may lead to dysregulated Kv1 translocation at the membrane, leading to excessive K currents. Of note, both Kv1 channels and Syngap1 are developmentally regulated (Okaty et al., 2009; Gamache et al., 2020), thus Syngap1 hypofunction may lead to dysregulation of genes encoding voltage-gated potassium channel affecting cell maturation and excitability.

Revised version: Despite the differences observed in how Syngap1 haploinsufficiency affects the anatomy and physiological properties of the two PV+ cell subtypes, a shared deficit we observed was the decrease in intrinsic excitability, as suggested by the increased AP threshold affecting AP initiation. In PV+ cells, voltage-gated D-type K+ currents mediated by the Kv1 family strongly contribute to AP generation, making them effective targets for modifying AP latency, threshold and rheobase current in these cells (Wang et al., 1994; Goldberg et al., 2008; Zurita et al., 2018). Here, we indeed restored the excitability of mutant PV+ cells by pharmacologically inhibiting this channel family using α-DTX; however, whether Kv1 currents or/and channel density are altered in mutant PV+ cells remain to be investigated. Consistent with our findings, Arora et al. (2022) rescued pyramidal cell intrinsic excitability and neuronal morphology via lowering elevated potassium channels, by expressing a dominant negative form of the Kv4.2 potassium channel subunit, dnKv4.2, in developing Syngap1 mutant mice. Kv4.2 and Kv1 potassium channels modulate the intrinsic excitability of pyramidal cell and PV+ interneurons, respectively, indicating that the action of Syngap1 on potassium channels may be a general mechanism. A recent study focusing on the PSD interactomes of Syngap1 isolated from adult homogenized mouse cortex suggested a physical interaction between Syngap1 and the potassium channel auxiliary subunit Kvβ2 (Wilkinson et al., 2017). Since Kvβ2 regulates the translocation of Kv1 channels in dopaminergic neurons and potentially PV+ cells (Okaty et al., 2009; Yee et al., 2022), it is possible that Syngap1 haploinsufficiency may lead to dysregulated Kv1 translocation at the membrane, leading to excessive K currents. Of note, both Kv1 channels and Syngap1 are developmentally regulated (Okaty et al., 2009; Gamache et al., 2020), thus Syngap1 hypofunction may lead to dysregulation of genes encoding voltage-gated potassium channel affecting cell maturation and excitability.

Earlier version: Interestingly, in cHet SST+ cells, we did not observe deficits in intrinsic excitability, but we found a reduction in number of APs generated at different somatic current injection. The fact that we observed differences in spiking activity in absence of intrinsic excitability alterations, could be due to the heterogeneity of SST+ interneurons likely present in our dataset (Scala et al., 2019; Hostetler et al., 2023). In addition, it’s possible that other intrinsic factors, not assessed in this study, may have contributed to this effect. For example, in SST+ cells the differences in the AHP kinetics depended predominantly on the presence of a second slower AHP component impacting on the overall amplitude, slope and duration of the AHP (Riedemann et al., 2018). Recent studies also showed that SYNGAP1 interacts with Kv4 (Wilkinson et al., 2017). Further, Syngap1 haploinsufficiency leads to increased Kv4 function in pyramidal cells in somatosensory cortex (Arora et al. 2022). Since somato-dendritic Kv4 channels in SST+ interneurons contribute to the regulation of their firing (Serôdio and Rudy, 1998; Bourdeau et al., 2007), Syngap1 haploinsufficiency might affect SST+ cell excitability via this channel.

Revised version: Interestingly, the alterations in intrinsic excitability were less pronounced in cHet SST+ interneurons, primarily affecting their firing rate. This discrepancy could be due to the heterogeneity of SST+ interneurons likely present in our dataset (Scala et al., 2019; Hostetler et al., 2023). In addition, it’s possible that other intrinsic factors, not assessed in this study, may have contributed to this effect. For example, in SST+ cells the differences in the AHP kinetics depend predominantly on the presence of a second slower AHP component impacting the overall amplitude, slope and duration of the AHP (Riedemann et al., 2018). Recent studies also showed that SYNGAP1 interacts with Kv4 (Wilkinson et al., 2017). Since somato-dendritic Kv4 channels in SST+ interneurons contribute to the regulation of their firing (Serôdio and Rudy, 1998; Bourdeau et al., 2007), Syngap1 haploinsufficiency could affect SST+ cell excitability via this channel.

PV+ cells have a fundamental role in generating and maintaining gamma oscillations in the brain (Cardin et al., 2009). In particular, recent studies have also shown that deficiency in PV interneuron-mediated inhibition contribute to increased baseline cortical gamma rhythm (Spencer, 2012; Carlén et al., 2012), a phenotype we observed in the auditory cortex of germline Syngap1+/- mutant mice, SYNGAP1-ID patients (Carreño-Muñoz et al., 2022) and mice with conditional Syngap1 haploinsifficiency restricted to MGE-derived interneurons (Jadhav et al., 2024). Specifically, PV+ interneurons, which target the perisomatic domain of pyramidal neurons, are adapted for fast synchronization of network activity controlling spike timing of the excitatory network in auditory cortex (Wehr and Zador, 2003; Li et al. 2014). To sharpen the tuning of neighboring pyramidal cells, PV+ interneurons need to be more effectively recruited by excitatory inputs so that they can restrict the temporal summation of excitatory responses of their pyramidal cell targets and increase the temporal precision of their firing (Povysheva et al., 2006). Our study suggests that decrease in AMPA-mediated thalamocortical input onto LIV PV+ cells along with deficits in their intrinsic excitability could account for altered spike timing of pyramidal cells causing an increase in overall network excitability.

Earlier version: An open question remains whether synaptic properties differ among PV+ cell subtypes, which we could not address due to technical limitations. These include the lack of specific neurochemical markers to distinguish between the two PV+ subtypes (Ekins et al., 2020), and the use of a Cs+-based internal solution required for voltage-clamp experiments, which prevents the recording of neuronal firing. One possibility could be to correlate PV expression levels directly with AP half-width, since BC-short may express higher levels of PV compared to BC-broad as already found in the striatum using patch seq approach (Bengtsson Gonzales et al., 2020). Further, in our studies, we used a conditional mouse model where Syngap1 haploinsufficiency is restricted to specific cell types, namely MGE-derived interneurons. Since cell-type specific genetic mutations do not typically occur in humans, it would be interesting to investigate whether SYNGAP1-haploinsufficient human-derived neurons show alterations in specific GABAergic subpopulations-intrinsic and synaptic properties. Further, whether and how PV+ physiology is affected in global haploinsufficient mice remain to be addressed. Of note, global haploinsufficient Syngap1 mice, SYNGAP1-ID patients and MGE-restricted Syngap1 haploinsufficient mice show comparable abnormal phenotypes in cortical auditory processing (Carreño-Muñoz et al., 2022; Jadhav et al., 2024). Further experiments specifically targeting PV+ cell activity, using targeted chemogenetic or pharmacological approach (Kourdougli et al., 2023) are required to shed light on the role PV+ interneuron hypoactivity in these phenotypes.

Revised version: An open question remains whether synaptic properties differ among PV+ cell subtypes, which we could not address due to technical limitations. These include the lack of specific neurochemical markers to distinguish between the two PV+ subtypes (Ekins et al., 2020), and the use of a Cs+- based internal solution required for voltage-clamp experiments, which prevents the recording of neuronal firing. It could be interesting to correlate PV expression levels directly with AP half-width, since BC-short may express higher levels of PV compared to BC-broad as already found in the striatum using patch seq approach (Bengtsson Gonzales et al., 2020). Further, in our studies, we used a conditional mouse model where Syngap1 haploinsufficiency is restricted to specific cell types, namely MGE-derived interneurons. Since cell-type specific genetic mutations do not typically occur in humans, it would be interesting to investigate whether SYNGAP1- haploinsufficient human-derived neurons show alterations in specific GABAergic subpopulations- intrinsic and synaptic properties. Further, whether and how PV+ physiology is affected in global haploinsufficient mice remain to be addressed. Of note, global haploinsufficient Syngap1 mice, SYNGAP1-ID patients and MGE-restricted Syngap1 haploinsufficient mice show comparable abnormal phenotypes in cortical auditory processing (Carreño-Muñoz et al., 2022; Jadhav et al., 2024). Further experiments specifically targeting PV+ cell activity, using targeted chemogenetic or pharmacological approach (Kourdougli et al., 2023) are required to shed light on the role PV+ interneuron hypoactivity in these phenotypes.

Materials and methods

Mice

All procedures and experiments were done in accordance with the Comité Institutionnel de Bonnes Pratiques Animales en Recherche (CIBPAR) of the CHU Ste-Justine Research Center in line with the principles published in the Canadian Council on Animal’s Care’s. Mice were housed (2-5 per cage), maintained in a 12/12 h light/dark cycle, and given ad libitum access to food and water. Experiments were performed in 9-13 weeks-old male mice during the light phase. To investigate the effects of Syngap1 haploinsufficiency in cortical PV+ and SST+ interneurons, we generated mice heterozygous for the Syngap1 conditional allele (Syngap1f/f; Jackson Laboratories; #029303) under the Tg(Nkx2.1-Cre) driver line (Jackson Laboratories; #008661) and further crossed them with mice carrying the RCE allele (Jackson Laboratories; #032037) to generate Tg(Nkx2.1- Cre):RCEf/f:Syngap1+/+ and Tg(Nkx2.1-Cre):RCEf/f:Syngap1f/+mice, for targeted recordings. Nkx2.1-expressing cells were identified by the expression of EGFP, since the RCE allele allow Cre- dependent EGFP expression.

Acute slice preparation

Briefly, animals (age range, mean ± S.E.M.: 75.5 ± 1.8 postnatal days for control group and 72.1 ± 1.7 postnatal days in cHet group) were anaesthetized deeply with ketamine–xylazine (ketamine: 100 mg/kg, xylazine: 10 mg/kg), transcardially perfused with 25 mL of ice-cold cutting solution (containing the following in mM: 250 sucrose, 2 KCl, 1.25 NaH2PO4, 26 NaHCO3, 7 MgSO4, 0.5 CaCl2, and 10 glucose, pH 7.4, 330–340 mOsm/L) and decapitated. The brain was then dissected carefully and transferred rapidly into an ice-cold (0–4 °C) cutting solution. Auditory thalamocortical slices (thickness, 350 μm) containing A1 and the medial geniculate nucleus were prepared. For A1 slices, the cutting angle was 15 degrees from the horizontal plane (lateral raised; Cruikshank et al., 2002; Zhao et al., 2009; Meng et al., 2017). Auditory thalamocortical slices were cut in the previously mentioned ice-cold solution using a vibratome (VT1000S; Leica Microsystems or Microm; Fisher Scientific) and transferred to a heated (37.5 °C) oxygenated recovery solution containing the following (in mM): 124 NaCl, 2.5 KCl, 1.25 NaH2PO4, 26 NaHCO3, 3 MgSO4, 1 CaCl2, and 10 glucose; pH 7.4; 300 mOsm/L, and allowed to recover for 45 min. Subsequently, during experiments, slices were continuously perfused (2 mL/min) with standard artificial cerebrospinal fluid (ACSF) containing the following (in mM): 124 NaCl, 2.5 KCl, 1.25 NaH2PO4, 26 NaHCO3, 2 MgSO4, 2CaCl2, and 10 glucose, pH 7.4 saturated with 95% O2 and 5% CO2 at near physiological temperature (30–33°C). We did not observe any difference between control and cHet mice in terms of slices quality, success rate of recordings and cellular health.

Whole-cell patch clamp recording

Earlier version: PV+ and SST+ neurons located in layer IV of A1 cortex were visually identified as EGFP-expressing somata under an epifluorescence microscope with blue light (filter set: 450–490 nm). All electrophysiological recordings were carried out using a 40x water-immersion objective. Recording pipettes were pulled from borosilicate glass (World Precision Instruments) with a PP-83 two-stage puller (Narishige) to a resistance range of 5–7 MΩ when backfilled with intracellular solution. Whole-cell patch-clamp recordings from PV+ and SST+ interneurons were performed in voltage or current-clamp mode. Pipette capacitance was neutralized, and bridge balance applied. For voltage-clamp recording, we used an intracellular Cs+-based solution containing (in mM): 130 CsMeSO4, 5 CsCl, 2 MgCl2, 10 phosphocreatine, 10 HEPES, 0.5 EGTA, 4 ATP-TRIS, 0.4 GTP-TRIS, 0.3% biocytin, 2 QX-314 (pH 7.2–7.3; 280–290 mOsm/L). These recordings were performed to analyze the excitatory drive received by PV+ and SST+ cells. Series resistance in voltage-clamp were monitored throughout the experiment and cells that had substantial changes in series resistance (>15%) during recording were discarded. The reported voltage values were not compensated for the junction potential. Recordings of sEPSC, mEPSCs were performed in voltage-clamp at –70 mV in the presence of gabazine (1μM; Tocris Bioscience) and CGP55845 (2μM; Abcam Biochemicals) for sEPSCs and with addition of tetrodotoxin for mEPSCs (TTX; 1μM; Alomone Labs). Recordings of spontaneous and miniature inhibitory postsynaptic currents (sIPSC, mIPSCs) were performed in voltage-clamp at +10mV in the presence of NBQX (10μM; Abcam Biochemicals) and DL-AP5 (100μM; Abcam Biochemicals) for sISPCs and with addition of TTX (1μM) for mIPSCs. For recordings of thalamocortical electrically evoked AMPA/NMDA EPSC in layer IV PV+ cells, current pulses (2ms duration, 25 to 1000 µA) were delivered to the thalamic radiation every 30 seconds. via a tungsten concentric bipolar microelectrode placed in the white matter midway between the medial geniculate nucleus and the AI (rostral to the hippocampus). Electrical stimulation of thalamic radiation may activate not only monosynaptic thalamic fibers but also polysynaptic (corticothalamic and/or corticocortical) EPSC component. To identify monosynaptic thalamo-cortical connections, we used as criteria the onset latencies of EPSC and the variability jitter obtained from the standard deviation of onset latencies. Onset latencies were defined as the time interval between the beginning of the stimulation artifact and the onset of the EPSC. Monosynaptic connections are characterized by short onset latencies and low jitter variability (Richardson et al., 2009; Blundon et al., 2011; Chun et al., 2013). In our experiments, the initial slopes of EPSCs evoked by white matter stimulation had short onset latencies (mean onset latency, 4.27 ± 0.11 ms, N=16 neurons in controls, and 5.07 ± 0.07 ms, N=14 neurons in cHet mice) and low onset latency variability jitter (0.24 ± 0.03 ms in controls vs 0.31 ± 0.03 ms in cHet mice), suggestive of activation of monosynaptic thalamocortical monosynaptic connections (Richardson et al., 2009; Blundon et al., 2011; Chun et al., 2013). Evoked AMPA (eAMPA) currents were recorded at -70 mV in presence of CGP (2μM) and gabazine (1μM), while evoked NMDA (eNMDA) currents were recorded at +40 mV in presence of NBQX (10µM) and confirmed after with the application of DL-AP5 (100µM). For paired pulse ratio (PPR) experiments, local synaptic stimulation in layer IV was achieved using a bipolar stimulating electrode made from borosilicate theta-glass capillaries (BT-150-10, Sutter Instruments, Novato, CA) filled with ACSF. Two EPSCs were evoked in layer IV BC with two electric pulses (0.2 ms, 40–140 µA each) at five intervals (30, 50, 100, 200 and 500 ms) every 10 secs in presence of picrotoxin (100 µM; Tocris/Cedarlane).

Revised version: PV+ and SST+ neurons located in layer IV of A1 cortex were visually identified as EGFP- expressing somata under an epifluorescence microscope with blue light (filter set: 450–490 nm). All electrophysiological recordings were carried out using a 40x water-immersion objective. Recording pipettes were pulled from borosilicate glass (World Precision Instruments) with a PP-83 two-stage puller (Narishige) to a resistance range of 5–7 MΩ when backfilled with intracellular solution. Whole-cell patch-clamp recordings from PV+ and SST+ interneurons were performed in voltage or current-clamp mode. Pipette capacitance was neutralized, and bridge balance applied. For voltage- clamp recording, we used an intracellular Cs+-based solution containing (in mM): 130 CsMeSO4, 5 CsCl, 2 MgCl2, 10 phosphocreatine, 10 HEPES, 0.5 EGTA, 4 ATP-TRIS, 0.4 GTP-TRIS, 0.3% biocytin, 2 QX-314 (pH 7.2–7.3; 280–290 mOsm/L). These recordings were performed to analyze the excitatory drive received by PV+ and SST+ cells. Series resistance in voltage-clamp were monitored throughout the experiment and cells that had substantial changes in series resistance (>15%) during recording were discarded. The reported voltage values were not compensated for the junction potential. Recordings of sEPSC and mEPSCs were performed in voltage-clamp at –70 mV in the presence of gabazine (1μM; Tocris Bioscience) and CGP55845 (2μM; Abcam Biochemicals). sEPSCs were first sampled over a 2 min period. A TTX (1μM; Alomone Labs) containing perfusion solution was then added (flow rate of 2 mL/min) and, after a 5 min-interval, mEPSCs were sampled over a 2 min period. Recordings of spontaneous and miniature inhibitory postsynaptic currents (sIPSC, mIPSCs) were performed in voltage-clamp at +10mV in the presence of NBQX (10μM; Abcam Biochemicals) and DL-AP5 (100μM; Abcam Biochemicals) for sISPCs and with addition of TTX (1μM) for mIPSCs. For recordings of thalamocortical electrically evoked AMPA/NMDA EPSC in layer IV PV+ cells, current pulses (2ms duration, 25 to 1000 µA) were delivered to the thalamic radiation every 30 seconds. via a tungsten concentric bipolar microelectrode placed in the white matter midway between the medial geniculate nucleus and the AI (rostral to the hippocampus). Electrical stimulation of thalamic radiation may activate not only monosynaptic thalamic fibers but also polysynaptic (corticothalamic and/or corticocortical) EPSC component. To identify monosynaptic thalamo-cortical connections, we used as criteria the onset latencies of EPSC and the variability jitter obtained from the standard deviation of onset latencies. Onset latencies were defined as the time interval between the beginning of the stimulation artifact and the onset of the EPSC. Monosynaptic connections are characterized by short onset latencies and low jitter variability (Richardson et al., 2009; Blundon et al., 2011; Chun et al., 2013). In our experiments, the initial slopes of EPSCs evoked by white matter stimulation had short onset latencies (mean onset latency, 4.27 ± 0.11 ms, N=16 neurons in controls, and 5.07 ± 0.07 ms, N=14 neurons in cHet mice) and low onset latency variability jitter (0.24 ± 0.03 ms in controls vs 0.31 ± 0.03 ms in cHet mice), suggestive of activation of monosynaptic thalamocortical monosynaptic connections (Richardson et al., 2009; Blundon et al., 2011; Chun et al., 2013). Evoked AMPA (eAMPA) currents were recorded at -70 mV in presence of CGP (2μM) and gabazine (1μM), while evoked NMDA (eNMDA) currents were recorded at +40 mV in presence of NBQX (10µM) and confirmed after with the application of DL-AP5 (100µM). For paired pulse ratio (PPR) experiments, local synaptic stimulation in layer IV was achieved using a bipolar stimulating electrode made from borosilicate theta-glass capillaries (BT-150-10, Sutter Instruments, Novato, CA) filled with ACSF. Two EPSCs were evoked in layer IV BC with two electric pulses (0.2 ms, 40–140 µA each) at five intervals (30, 50, 100, 200 and 500 ms) every 10 secs in presence of picrotoxin (100 µM; Tocris/Cedarlane).

Earlier version: Current-clamp recordings were obtained in ACSF containing synaptic blockers gabazine (1μM), CGP55845 (2μM) and kynurenic acid (2mM). For these recordings, we used an intracellular K+-based solution containing (in mM): 130 KMeSO4, 2 MgCl2, 10 di-Na-phosphocreatine, 10 HEPES, 4 ATP-Tris, 0.4 GTP-Tris, and 0.3% biocytin (Sigma), pH 7.2–7.3, 280–290 mOsm/L. Passive and active membrane properties were analyzed in current clamp mode: active membrane properties were recorded by subjecting cells to multiple current step injections (step size 40 pA) of varying amplitudes (–200 to 600 pA). In subsequent experiments, the same protocol was repeated in presence of the voltage gated potassium channel (Kv) blocker α-DTX (100nM; Alomone labs). In these experiments, slices were recovered in a holding chamber for at least 1 hr in presence of α-DTX. Once placed in recording chamber slices were kept for recording for max 1 hr. Passive membrane properties, resting membrane potential (Vm), input resistance (Rin), and membrane capacitance (Cm) were obtained immediately after membrane rupture. Membrane potentials were maintained at –80 mV, series resistances (10–18 MΩ) and Rin were monitored on-line with a 40 pA current injection (150ms) given before each 500ms current injection stimulus. Only cells with resting membrane potential more negative than -60mV at the start of recording and spikes with overshoot were considered for further analysis. We performed bridge balance and neutralized the capacitance before starting every recording. The bridge balance was monitored throughout the experiment, and neurons showing changes of >15% in bridge balance during the recording were discarded. Data acquisition (filtered at 2–3 kHz and digitized at 10kHz; Digidata 1440, Molecular Devices, CA, United States) was performed using the Multiclamp 700B amplifier and the Clampex 10.6 software (Molecular Devices).

Revised version: Current-clamp recordings were obtained in ACSF containing synaptic blockers gabazine (1μM), CGP55845 (2μM) and kynurenic acid (2mM). For these recordings, we used an intracellular K+- based solution containing (in mM): 130 KMeSO4, 2 MgCl2, 10 di-Na-phosphocreatine, 10 HEPES, 4 ATP-Tris, 0.4 GTP-Tris, and 0.3% biocytin (Sigma), pH 7.2–7.3, 280–290 mOsm/L. Passive and active membrane properties were analyzed in current clamp mode: active membrane properties were recorded by subjecting cells to multiple current step injections (step size 40 pA) of varying amplitudes (–200 to 600 pA). In subsequent experiments, the same protocol was repeated in presence of the voltage gated potassium channel (Kv) blocker α-DTX (100nM; Alomone labs). In these experiments, slices were recovered in a holding chamber for at least 1 hr in presence of α- DTX. Once placed in recording chamber slices were kept for recording for max 1 hr. Passive membrane properties, resting membrane potential (Vm), input resistance (Rin), and membrane capacitance (Cm) were obtained immediately after membrane rupture. Membrane potentials were maintained at –80 mV, series resistances (10–18 MΩ) and Rin were monitored on-line with a 40 pA current injection (150ms) given before each 500ms current injection stimulus. Only cells with resting membrane potential more negative than -60mV at the start of recording and spikes with overshoot were considered for further analysis. We performed bridge balance and neutralized the capacitance before starting every recording. The bridge balance was monitored throughout the experiment, and neurons showing changes of >15% in bridge balance during the recording were discarded. Data acquisition (filtered at 2–3 kHz and digitized at 10kHz; Digidata 1440, Molecular Devices, CA, United States) was performed using the Multiclamp 700B amplifier and the Clampe× 10.6 software (Molecular Devices).

Electrophysiological data analysis

All analysis was performed by researchers blinds to the genotype. Analysis of electrophysiological recordings was performed using Clampfit 10.7 (Molecular Devices). For the analysis of sEPSCs, mEPSCs, sIPSCs and mIPSCs a minimum of 100 events were sampled per cell over a 2min period using an automated template search algorithm in Clampfit. The 20–80% rise time of the response and the decay time constant determined from the exponential fit (100–37%) were calculated. Charge transfer was calculated by integrating the area under the EPSC and IPSC waveforms. The mean PSC synaptic current was calculated as the charge transfer of the averaged PSC (ΔQ) multiplied by mean PSC frequency. For thalamocortical eAMPA and eNMDA, the mean amplitude of EPSCs including both failure and success was obtained from a total of 5 to 10 sweeps. Onset latency indicated the time from beginning of stimulus artifact and the onset of eAMPA or eNMDA. To measure eNMDA/eAMPA ratios, the eAMPA component was taken at the peak of EPSC at −70 mV, whereas the eNMDA component was measured at the peak of EPSC at +40 mV. For PPR experiments, the PPR was calculated as the ratio between the mean AUC of the second response and the mean AUC of the first response (7 to 10 sweeps).

Earlier version: For the analysis of current-clamp recordings from PV cells, rheobase was measured as the minimal current necessary to evoke an AP. For the analysis of the AP properties, the first AP appearing within a 50ms time window from beginning of current pulse was analyzed. AP latency was measured as the time between current step onset and when membrane voltage reached AP threshold. The AP amplitude was measured from the AP threshold to the peak. The AP half-width was measured at the voltage level of the half of AP amplitude. The AP rise and fall time were measured between the AP threshold and the maximal AP amplitude, and between the maximal AP amplitude and the AP end, respectively. The fast afterhyperpolarization (fAHP) amplitude was determined as the minimum voltage following the AP peak subtracted from the AP threshold. fAHP time was determined as the time between AP threshold and the negative peak of fAHP. The hyperpolarization-activated cation current (Ih)-associated voltage rectification (Ih sag) was determined as the amplitude of the membrane potential sag from the peak hyperpolarized level to the level at the end of the hyperpolarizing step when the cell was hyperpolarized to –100mV. Membrane time constant (τ) was calculated by the product of Rin and Cm. For the firing analysis, we considered only APs with amplitude >30 mV as full APs. Spikelets with amplitude smaller than 30mV were not analyzed in this study. The inter-spike interval (ISI) was determined by the time difference between adjacent AP peaks. Spike amplitude accommodation ratio (AAR) was calculated by dividing the amplitude of the last AP by the amplitude of the first generated in response to 2x rheobase current injection. Firing frequency adaptation ratios (FFAR) were calculated by dividing the last ISI with the first one of the responses to 2x rheobase current injection. The maximal (initial) firing frequency (Fmaxinitial) was computed as the reciprocal of the average of the first 2 ISIs in a spike train elicited by the current step (max +600pA) applied before a noticeable appearance of spikelets. The steady-state firing frequency (Fss) was computed as the reciprocal of the average of the last 4 ISIs in the spike train were Fmaxinitial was obtained. Finally, the number of APs (# APs)-current relationship for evoked firing was determined by injecting 500-ms somatic current steps of increasing amplitude (40pA increments) to a maximum of 600pA. For current-clamp recordings in presence of α-DTX experiments the delta (Δ) values were calculated for rheobase, AP threshold and AP number at +200pA, by subtracting individual values of α-DTX-treated cells from the average of their respective control group.

Revised version: For the analysis of current-clamp recordings from PV cells, threshold current was measured as the minimal current necessary to evoke an AP. For the analysis of the AP properties, the first AP appearing within a 50ms time window from beginning of current pulse was analyzed. AP latency was measured as the time between current step onset and when membrane voltage reached AP threshold. The AP amplitude was measured from the AP threshold to the peak. The AP half-width was measured at the voltage level of the half of AP amplitude. The AP rise and fall time were measured between the AP threshold and the maximal AP amplitude, and between the maximal AP amplitude and the AP end, respectively. The fast afterhyperpolarization (fAHP) amplitude was determined as the minimum voltage following the AP peak subtracted from the AP threshold. fAHP time was determined as the time between AP threshold and the negative peak of fAHP. The hyperpolarization-activated cation current (Ih)-associated voltage rectification (Ih sag) was determined as the amplitude of the membrane potential sag from the peak hyperpolarized level to the level at the end of the hyperpolarizing step when the cell was hyperpolarized to –100mV. Membrane time constant (τ) was calculated by the product of Rin and Cm. For the firing analysis, we considered only APs with amplitude >30 mV as full APs. Spikelets with amplitude smaller than 30mV were not analyzed in this study. The inter-spike interval (ISI) was determined by the time difference between adjacent AP peaks. Spike amplitude accommodation ratio (AAR) was calculated by dividing the amplitude of the last AP by the amplitude of the first generated in response to 2x threshold current injection. Firing frequency adaptation ratios (FFAR) were calculated by dividing the last ISI with the first one of the responses to 2x threshold current injection. The maximal (initial) firing frequency (Fmaxinitial) was computed as the reciprocal of the first ISI in a spike train elicited by the current step (max +600pA) applied before a noticeable appearance of spikelets. The steady-state firing frequency (Fss) was computed as the reciprocal of the average of the last 4 ISIs in the spike train were Fmaxinitial was obtained. Finally, the number of APs (# APs)-current relationship for evoked firing was determined by injecting 500-ms somatic current steps of increasing amplitude (40pA increments) to a maximum of 600pA. For current-clamp recordings in presence of α-DTX experiments the delta (Δ) values were calculated for threshold current, AP threshold and AP number at +200pA, by subtracting individual values of α-DTX-treated cells from the average of their respective control group.

Hierarchical clustering and principal component analysis (PCA)

Earlier version: This analysis was based on previous published data finding heterogeneity in PV+ interneurons population (Helm et al., 2013), and performed using the software IBM SPSS V29.0.0. Hierarchical clustering was based on Euclidean distance of PV+ cells from control mice. To identify potential clusters, we used AP half-width and Fmaxinitial values. We then performed multidimensional cluster analysis on passive and active membrane properties to identify possible common groupings of PV+ interneurons using the software Prism 9.0 (GraphPad Software). In our database, we focused on 13 parameters (see Figure 6b) that were for the majority unrelated. Figure 6 is a cross-correlation matrix of these 13 parameters with correlation indices shade-coded. Black means perfectly positively correlated (correlation index of 1.0), white means perfectly negative correlated (correlation index of -1.0) and light gray not correlated (correlation index of 0). In our database, we have parameters that are not strongly correlated (e.g., rheobase and AP amplitude, fAHP amplitude and AP latency), and others that are correlated (correlation coefficient > 0.5 or < -0.5; AP latency and amplitude AR; AP half-width and AP amplitude; Rin and AP half-width; AP amplitude and Fmax initial). We retained all parameters because they encompass different features of membrane properties (Helm et al., 2013). We therefore performed PCA on the 13 parameters to reduce the dimensionality and to potentiate clusters separation.

Revised version: This analysis was based on previous published data finding heterogeneity in PV+ interneurons population (Helm et al., 2013), and performed using the software IBM SPSS V29.0.0. Hierarchical clustering was based on Euclidean distance of PV+ cells from control mice. To identify potential clusters, we used AP half-width and Fmaxinitial values. We then performed multidimensional cluster analysis on passive and active membrane properties to identify possible common groupings of PV+ interneurons using the software Prism 9.0 (GraphPad Software). In our database, we focused on 13 parameters (see Figure 5b) that were for the majority unrelated. Figure 5b is a cross-correlation matrix of these 13 parameters with correlation indices shade-coded. Black means perfectly positively correlated (correlation index of 1.0), white means perfectly negative correlated (correlation index of -1.0) and light gray not correlated (correlation index of 0). In our database, we have parameters that are not strongly correlated (e.g., threshold current and AP amplitude, fAHP amplitude and AP latency), and others that are correlated (correlation coefficient > 0.5 or < -0.5; AP latency and amplitude AR; AP half-width and AP amplitude; Rin and AP half-width; AP amplitude and Fmax initial). We retained all parameters because they encompass different features of membrane properties (Helm et al., 2013). We therefore performed PCA on the 13 parameters to reduce the dimensionality and to potentiate clusters separation.

Immunohistochemistry, cell reconstruction and anatomical identification

Earlier version: For post-hoc anatomical identification, every recorded neuron was filled with biocytin (0.5%, Sigma) during whole-cell recordings (15 min). To reveal biocytin, the slices were permeabilized with 0.3% triton X-100 and incubated at 4 °C with a streptavidin-conjugated Alexa-488 (1:1000, Invitrogen, Cat# S11223) in TBS. For PV and SST immunofluorescence, sections were permeabilized with 0.25% Triton X-100 in PBS and incubated in blocking solution containing 20% normal goat serum (NGS) for 1 h. Then, sections were incubated with the following primary antibodies diluted in 1% NGS, 0.25% Triton-X 100 in PBS: were incubated with primary antibodies mouse anti-PV (1:1000, Swant, Cat# 235) and rabbit anti-SST (1:1000, Thermofisher Invitrogen, Cat# PA5-82678) at 4°C for 48-72 h. Sections were then washed in PBS (3X-10 mins each), incubated for 2 hrs at RT with the following secondary antibodies diluted in 1% NGS, 0.25% Triton-X 100 in PBS and mounted on microscope slides: Alexa 555-conjugated goat anti-rabbit (1:1000; Life technologies, A21430) and Alexa 647-conjugated goat anti-mouse (1:250, Cell signaling, 4410S). Z-stacks of biocytin-filled cells were acquired with a 1μm step using a 20x objective (NA 0.75) on the Leica SP8-DLS confocal microscope. Confocal stacks of PV+ neurons were merged for detailed reconstruction in Neuromantic tracing software version 1.7.5 (Myatt et al., 2012). Dendritic arbors were reconstructed plane-by-plane from the image z-stack and analyzed using the Neuromantic software. All the reconstructions used for dendritic analysis contained intact cells with no evident cut dendrites. Sholl analysis of reconstructed dendritic arbors was performed in FIJI software using the plugin Neuroanatomy. This analysis was performed in radial coordinates, using a 10μm step size from r=0, with the origin centered on the cell soma, and counting the number of compartments crossing a given radius. All analysis was performed by researchers blinds to the genotype.

Revised version: For post-hoc anatomical identification, every recorded neuron was filled with biocytin (0.5%, Sigma) during whole-cell recordings (15 min). To reveal biocytin, the slices were permeabilized with 0.3% triton X-100 and incubated at 4 °C with a streptavidin-conjugated Alexa-488 (1:1000, Invitrogen, Cat# S11223) in TBS. For PV and SST immunofluorescence, sections were permeabilized with 0.25% Triton X-100 in PBS and incubated in blocking solution containing 20% normal goat serum (NGS) for 1 h. Then, sections were incubated with the following primary antibodies diluted in 1% NGS, 0.25% Triton-X 100 in PBS: were incubated with primary antibodies mouse anti-PV (1:1000, Swant, Cat# 235) and rabbit anti-SST (1:1000, Thermofisher Invitrogen, Cat# PA5-82678) at 4°C for 48-72 h. Sections were then washed in PBS (3X- 10 mins each), incubated for 2 hrs at RT with the following secondary antibodies diluted in 1% NGS, 0.25% Triton-X 100 in PBS and mounted on microscope slides: Alexa 555-conjugated goat anti-rabbit (1:1000; Life technologies, A21430) and Alexa 647-conjugated goat anti-mouse (1:250, Cell signaling, 4410S). Whole biocytin-filled cells were acquired with a 1μm step using a 20x objective (NA 0.75) on a Leica SP8-DLS confocal microscope. For each imaged neuron, we noted the spatial distribution across cortical layers of the axonal arbor and whether its dendrites carried spines. Confocal stacks of PV+ neurons were merged for detailed reconstruction in Neuromantic tracing software version 1.7.5 (Myatt et al., 2012). Dendritic arbors were reconstructed plane-by-plane from the image z-stack and analyzed using the Neuromantic software. All the reconstructions used for dendritic analysis contained intact cells with no evident cut dendrites. Sholl analysis of reconstructed dendritic arbors was performed in FIJI software using the plugin Neuroanatomy. This analysis was performed in radial coordinates, using a 10μm step size from r=0, with the origin centered on the cell soma, and counting the number of compartments crossing a given radius. All analysis was performed by researchers blinds to the genotype.

vGlut1/PSD95 and vGlut 2/PSD immunostaining, imaging and quantification

P60 mice were anesthetized with: Ketamine-100mg/kg+Xylazine-10mg/kg+Acepromazine- 10mg/kg and perfused transcardially with 0.9% saline followed by 4% Paraformaldehyde (PFA) in phosphate buffer (0.1M PB, pH 7.2-7.4). Brains were dissected out and post-fixed in 4% PFA overnight at 4°C. They were subsequently transferred to 30% sucrose (prepared in PBS, pH 7.2) at 4°C for 48 hrs. Brains were then embedded in molds filled with OCT Tissue Tek and frozen in a bath of 2-methybutane placed on a bed of dry ice and ethanol. Coronal sections were cut at 40μm with a cryostat (Leica CM3050 S) and collected as floating sections in PBS. Brain sections were first permeabilized in 0.2 % Triton-X in PBS for 1 hr at RT and then blocked in 10% normal donkey serum (NDS) with 0.2% Triton-X 100 and 5% bovine serum albumin (BSA) in PBS for 2 hrs at RT followed by incubation at 4°C for 48 hrs with the following primary antibodies diluted in 5% NDS, 0.2% Triton-X 100 and 2 % BSA in PBS: goat anti-PV (1:1000, Swant, Cat# PVG-213), rabbit anti- VGlut1 (1:100 Thermo Fisher/ Invitrogen, Cat# 48-2400), rabbit anti-VGlut2 1:1000, Synaptic systems, Cat# 135402 ) mouse anti-PSD95 (1:500, Invitrogen, Cat# MA1-05). Sections were then washed in PBS + 0.1% Triton-X 100 (3X10 mins each) and incubated for 2 hrs at RT with the following secondary antibodies diluted in 5% NDS, 0.2% Triton-X 100 and 2 % BSA in PBS : Alexa 488-conjugated donkey anti-rabbit (1:500, Life technologies/ Invitrogen, Cat# A11055 ), Alexa 555-conjugated donkey anti-mouse (1:500, Life technologies/ Invitrogen, Cat# A31570), Alexa 633-conjugated donkey anti-goat (1:500, Invitrogen, Cat# A21082). Sections were rinsed in 0.1% Triton-X 100 in PBS (3X10’ each +1X5’) and mounted with Vectorshield mounting medium (Vector).

Earlier version: Confocal Imaging and quantification: Immunostained sections were imaged using a Leica SP8-STED confocal microscope, with an oil immersion 63x (NA 1.4) at 1024 X 1024, zoom=1, z-step =0.3 μm, stack size of ∼15 μm. Images were acquired from the A1 from at least 3 coronal sections per animal. All the confocal parameters were maintained constant throughout the acquisition of an experiment. All images shown in Supplemental Figure 2 a,c are from a single confocal plane. To quantify the number of vGlut1/PSD95 or vGlut2/PSD95 putative synapses, images were exported as TIFF files and analyzed using Fiji (Image J) software. We first manually outlined the profile of each PV cell soma (identified by PV immunolabeling). At least 4 innervated somata were selected in each confocal stack. We then used a series of custom-made macros in Fiji as previously described (Chehrazi et al, 2023). After subtracting background (rolling value = 10) and Gaussian blur (σ value = 2) filters, the stacks were binarized and vGlut1/PSD95 or vGlut2/PSD95 puncta were independently identified around the perimeter of a targeted soma in the focal plane with the highest soma circumference. Puncta were quantified after filtering particles for size (included between 0-2μm2) and circularity (included between 0-1). Quantification of the density of perisomatic puncta colocalizing both VGlut1-PSD95 and VGlut2-PSD95 were normalised to controls. All analysis was performed by researchers blinds to the genotype. No mouse was excluded from this analysis.

Revised version: Immunostained sections were imaged using a Leica SP8-STED confocal microscope, with an oil immersion 63x (NA 1.4) at 1024 X 1024, zoom=1, z-step =0.3 μm, stack size of ∼15 μm. Images were acquired from the A1 from at least 3 coronal sections per animal. All the confocal parameters were maintained constant throughout the acquisition of an experiment. All images shown in Suppl. Figure 3a,c are from a single confocal plane. To quantify the number of vGlut1/PSD95 or vGlut2/PSD95 putative synapses, images were exported as TIFF files and analyzed using Fiji (Image J) software. We first manually outlined the profile of each PV cell soma (identified by PV immunolabeling). At least 4 innervated somata were selected in each confocal stack. We then used a series of custom-made macros in Fiji as previously described (Chehrazi et al, 2023). After subtracting background (rolling value = 10) and Gaussian blur (σ value = 2) filters, the stacks were binarized and vGlut1/PSD95 or vGlut2/PSD95 puncta were independently identified around the perimeter of a targeted soma in the focal plane with the highest soma circumference. Puncta were quantified after filtering particles for size (included between 0-2μm2) and circularity (included between 0-1). Quantification of the density of perisomatic puncta colocalizing both VGlut1-PSD95 and VGlut2-PSD95 were normalised to controls. All analysis was performed by researchers blinds to the genotype. No mouse was excluded from this analysis.

Statistics

Earlier version: Data were expressed as mean±SEM. For statistical analysis, we based our conclusion on the statistical results generated by linear mixed model (LMM), modelling animal as a random effect and genotype as fixed effect. We used this statistical analysis because we considered the number of mice as independent replicates and the number of cells in each mouse as repeated measures (Berryer et al. 2016; Heggland et al., 2019; Yu et al., 2022). In parallel, we also tested the data for normality with a Shapiro–Wilcoxon test, then if data were normally distributed, standard parametric statistics were used (unpaired t-test), while if data were not normally distributed, non-parametric Mann–Whitney test was used for comparisons of two groups. The two statistical analyses were always performed for all the data of the study and indicated in the tables related to each figure (Tables 1-8, and S1-S2). Two-way ANOVA with Sidak’s multiple comparison post hoc test was used for the detection of differences in AP firing and number of dendritic intersections between genotypes. Statistical analysis was performed using Sigma Plot 11.0, Prism 9.0 (GraphPad Software) and IBM SPSS V29.0.0 for LMM analysis.

Revised version: Data were expressed as mean±SEM. For statistical analysis, we based our conclusion on the statistical results generated by LMM, modelling animal as a random effect and genotype as fixed effect. We used this statistical analysis because we considered the number of mice as independent replicates and the number of cells in each mouse as repeated measures (Berryer et al. 2016; Heggland et al., 2019; Yu et al., 2022). For cumulative distributions, the same number of events was chosen randomly from each cell and analysed by LMM, modelling animal as a random effect and genotype as fixed effect. Two-way ANOVA with Sidak’s multiple comparison post hoc test was used for the detection of differences in AP firing and number of dendritic intersections between genotypes. Statistical analysis was performed using Sigma Plot 11.0, Prism 9.0 (GraphPad Software) and IBM SPSS V29.0.0 for LMM analysis.